---
title: "Topic 5: Discrete Probability Distributions (Spreadsheets, Dynamic)"
format: 
  live-html:
    include-in-header:
      text: |
        <style>
        .btn-exercise-solution { display: none !important; }
        </style>
    webr:
      packages:
        - gradethis
        - jsonlite
        - base64enc
      cell-options: 
        persist: true
    live:
      show-solutions: true
---

{{< include ./_extensions/r-wasm/live/_knitr.qmd >}}
{{< include ./_extensions/r-wasm/live/_gradethis.qmd >}}

```{webr}
#| edit: false
#| output: false
#| define:
#|   - ok_response

library(htmltools)
ok_response <- function(response, n) {
  if (is.na(response)) div(HTML("You haven't answered yet."), style = "color: purple")
  else if (response == n) div(HTML("Correct ✓"), style = "color: green")
  else div(HTML("Incorrect ✗"), style = "color: red")
}
```

```{ojs}
//| echo: false

function ok_checkbox(response, n) {
  if (!response || response.length === 0)
    return html`<span style="color:purple">You haven't answered yet.</span>`;
  if (response.toString() === n)
    return html`<span style="color:green">Correct ✓</span>`;
  return html`<span style="color:red">Not Yet! ✗</span>`;
}
```

```{webr}
#| edit: false
#| output: false
#| define:
#|   - normalize_checkbox_response

normalize_checkbox_response <- function(x) {
  return(cat(sort(x), sep = ","))
}
```

```{webr}
#| context: setup
#| echo: false
#| output: false
#| define:
#|   - n_students
#|   - p_assign
#|   - p_pct
#|   - thresh_atmost
#|   - thresh_atleast
#|   - thresh_between_lo
#|   - thresh_between_hi

library(jsonlite)
library(base64enc)

set.seed(NULL)

# Sample size: integer between 15 and 25
n_students <- sample(15:25, 1)

# Probability: one of several "clean" percentage values
p_assign <- sample(c(0.25, 0.28, 0.30, 0.32, 0.34, 0.36, 0.38, 0.40), 1)
p_pct <- round(p_assign * 100)

# Thresholds derived from n so they are always pedagogically sensible
thresh_atmost <- round(n_students * p_assign) + sample(1:3, 1)
thresh_atleast <- round(n_students * p_assign) + sample(4:7, 1)
thresh_atleast <- min(thresh_atleast, n_students - 1)
thresh_between_lo <- max(2, round(n_students * p_assign) - sample(2:4, 1))
thresh_between_hi <- min(n_students - 1, round(n_students * p_assign) + sample(3:6, 1))

# Initialise exercise results
exercise_results <- list(
  notebook              = "Topic 5: Discrete Probability Distributions",
  prob_a_and_b          = "unattempted",
  chem_quiz_binom       = "unattempted",
  binom_assignments_ex1 = "unattempted",
  binom_assignments_ex2 = "unattempted",
  binom_assignments_ex3 = "unattempted",
  binom_assignments_ex4 = "unattempted",
  binom_assignments_ex5 = "unattempted",
  binom_assignments_ex6 = "unattempted",
  binom_assignments_ex7 = "unattempted"
)
```

```{r setup, include=FALSE}
library(gradethis)
```

:::{.callout-note}
## About

This activity provides an introduction to discrete probability, including basic probability through counting outcomes, and calculating probabilities associated with outcomes of binomial experiments using the binomial distribution.

```{ojs}
//| echo: false

{
  const btn = html`<button style="
    background: #436b95;
    color: white;
    border: none;
    padding: 8px 16px;
    border-radius: 4px;
    cursor: pointer;
    font-size: 0.9em;
  ">🔄 Reset My Responses</button>`;

  btn.onclick = () => {
    if (confirm('Reset all responses for this activity? This cannot be undone.')) {
      Object.keys(localStorage)
        .filter(k => k.endsWith('_selected_t5d') || k.endsWith('_correct_t5d') || k.endsWith('_result_t5d'))
        .forEach(k => localStorage.removeItem(k));
      location.reload();
    }
  };

  return btn;
}
```

*Note.* The button above resets multiple choice and checkbox questions. Currently, resetting answer entry cells must be done manually via hitting the *Start Over* button on each individual interactive cell.

:::

:::{.callout-note}
## Using a Spreadsheet Program
There are times throughout this activity where you'll be asked to do work inside of your favorite spreadsheet program. 

+ If you are a Google Sheets user, [open a blank Google Sheet](https://sheets.new){target="_blank"}
+ If you are an Excel user, open an empty workbook on your computer or [create a blank workbook in Excel online](https://office.live.com/start/Excel.aspx)
:::

## Discrete Probability Distributions

Throughout this activity you'll be introduced to the notion of probability and will explore applications of probability and discrete random variables. After developing some intuition using foundational probability ideas, we'll focus on binomial experiments and using the binomial distribution to find probabilities of prescribed outcomes.

:::{.callout-tip}
## Limitations

There are entire courses devoted to probability -- we will only cover probability to the extent that it is necessary for use in this course. If you are interested in a more detailed treatment of probability, seek out one of the many great courses available.
:::

## Objectives

**Activity Objectives:** After completing this workbook you should be able to:

+ Define, discuss, and interpret the probability of an event as its likelihood.
+ Apply fundamental counting principles and the notion of independence to compute the probability associated with the occurrence of a sequence of events.
+ Use the definition of binomial experiments to identify scenarios to which the binomial distribution can be applied.
+ Apply the binomial distribution in appropriate scenarios to find probabilities associated with specified outcomes using spreadsheet software.
+ Given a binomial experiment, compute the expected number of successful outcomes as well as the standard deviation for number of successes.

## Basic Probability

**Definition of Probability (*frequentist*):** For a given random process, the probability of an event $A$ is the proportion of time we would observe outcome $A$ if the random process were repeated an infinite number of times.

**Example:** Given a fair coin, the probability of a flip turning up heads is $0.5$ (or 50%). Similarly, given a fair six-sided die, the probability of a roll resulting in a number greater than four is $1/3$ (or about 33.3%) because there are two outcomes satisfying the criteria (rolling a 5 or rolling a 6) out of the six total possible outcomes.

**Try It!** Now it is your turn. Try the next few problems. Be sure to note any questions you have as you work through them.

:::{.callout-caution}
## Check Your Understanding: Basic Probability I

Given one fair, six-sided die, what is the probability of rolling a three?

```{ojs}
//| echo: false

mutable ok_response = (response, n) => { return html`Loading...` };
viewof q1 = Inputs.radio(
  new Map([
    ["0 / 6", 1],
    ["1 / 6", 2],
    ["2 / 6", 3],
    ["3 / 6", 4],
    ["4 / 6", 5],
    ["5 / 6", 6],
    ["6 / 6", 7]
  ]),
  {value: JSON.parse(localStorage.getItem("q1_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q1_selected_t5d", JSON.stringify(q1));
  localStorage.setItem("q1_correct_t5d", "2");
  localStorage.setItem("q1_result_t5d", q1 === null ? "unattempted" : (q1 == 2 ? "correct" : "incorrect"));
}
ok_response(q1, "2");
```

:::

:::{.callout-caution}
## Check Your Understanding: Basic Probability II

Given one fair, six-sided die, what is the probability of rolling a two, four, or six?

```{ojs}
//| echo: false

viewof q2 = Inputs.radio(
  new Map([
    ["0 / 6", 1],
    ["1 / 6", 2],
    ["2 / 6", 3],
    ["3 / 6", 4],
    ["4 / 6", 5],
    ["5 / 6", 6],
    ["6 / 6", 7]
  ]),
  {value: JSON.parse(localStorage.getItem("q2_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q2_selected_t5d", JSON.stringify(q2));
  localStorage.setItem("q2_correct_t5d", "4");
  localStorage.setItem("q2_result_t5d", q2 === null ? "unattempted" : (q2 == 4 ? "correct" : "incorrect"));
}
ok_response(q2, "4");
```

:::

:::{.callout-caution}
## Check Your Understanding: Basic Probability III

Given two fair, six-sided dice, which is larger?

```{ojs}
//| echo: false

viewof q3 = Inputs.radio(
  new Map([
    ["The probability of rolling a total of two.", 1],
    ["The probability of rolling a total of five.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q3_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q3_selected_t5d", JSON.stringify(q3));
  localStorage.setItem("q3_correct_t5d", "2");
  localStorage.setItem("q3_result_t5d", q3 === null ? "unattempted" : (q3 == 2 ? "correct" : "incorrect"));
}
ok_response(q3, "2");
```

:::

Good work on that last set of questions. In those problems you could find the probability by counting the number of ways the desired outcome could occur and then dividing that number by the total number of outcomes possible. In the last question, there were simply more ways to roll a five (four ways to do it) than to roll a two (just one way). What if we try doing something a bit more complicated? Say we wanted to know the probability of rolling at least a two on a single roll of a die *and then* flipping a "tails" on a single flip of a coin?

:::{.callout-tip}
## Probability and Independent Events
If $A$ and $B$ are independent events (that is, the probability that $B$ occurs does not depend on whether or not $A$ occurred, and vice-versa), then the probability of $A$ and $B$ occurring is the product of the probability of $A$ occurring and the probability of $B$ occurring. Mathematically, we write: $\mathbb{P}\left[A~\text{and}~B\right] = \mathbb{P}\left[A\right]\cdot\mathbb{P}\left[B\right]$.

:::

:::{.callout-caution}
## Check Your Understanding: Probability and Independent Events I

Given a single roll of a fair, six-sided die, what is the probability of rolling at least a two?

```{ojs}
//| echo: false

viewof q4 = Inputs.radio(
  new Map([
    ["0 / 6", 1],
    ["1 / 6", 2],
    ["2 / 6", 3],
    ["3 / 6", 4],
    ["4 / 6", 5],
    ["5 / 6", 6],
    ["6 / 6", 7]
  ]),
  {value: JSON.parse(localStorage.getItem("q4_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q4_selected_t5d", JSON.stringify(q4));
  localStorage.setItem("q4_correct_t5d", "6");
  localStorage.setItem("q4_result_t5d", q4 === null ? "unattempted" : (q4 == 6 ? "correct" : "incorrect"));
}
ok_response(q4, "6");
```

:::

:::{.callout-caution}
## Check Your Understanding: Probability and Independent Events II

Given a single flip of a fair coin, what is the probability of the coin landing with *tails* facing upwards?

```{ojs}
//| echo: false

viewof q5 = Inputs.radio(
  new Map([
    ["0 / 2", 1],
    ["1 / 2", 2],
    ["2 / 2", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q5_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q5_selected_t5d", JSON.stringify(q5));
  localStorage.setItem("q5_correct_t5d", "2");
  localStorage.setItem("q5_result_t5d", q5 === null ? "unattempted" : (q5 == 2 ? "correct" : "incorrect"));
}
ok_response(q5, "2");
```

:::

:::{.callout-caution}
## Check Your Understanding: Probability and Independent Events III

Use your spreadsheet to compute the probability that in a single roll of a fair die and a flip of a coin we observe a roll of at least two and a flip of tails. Enter your result below.

```{webr}
#| exercise: prob-a-and-b
# Enter your answer / result here

```

:::{.hint exercise="prob-a-and-b"}
:::{.callout-note collapse="false"}
## Hint 1

Use the probabilities you identified in the two questions above. You can multiply values together in a spreadsheet cell by typing `=` followed by the values separated by `*`. For example, typing `=(1/4)*(1/3)` and pressing Enter gives the product of those two fractions.

:::
:::

:::{.hint exercise="prob-a-and-b"}
:::{.callout-note collapse="false"}
## Hint 2

The flip of a coin and roll of a fair die are independent events, so we can multiply their individual probabilities together. In a blank spreadsheet cell, start with `=` and multiply the two probabilities.

:::
:::

:::{.hint exercise="prob-a-and-b"}
:::{.callout-note collapse="false"}
## Hint 3

In a blank cell, type `=(5/6)*(1/2)` and press Enter. The result is the probability of rolling at least a two **and** flipping tails.

:::
:::

```{webr}
#| exercise: prob-a-and-b
#| solution: true

(5/6)*(1/2)
```

```{webr}
#| exercise: prob-a-and-b
#| check: true

exercise_results$prob_a_and_b <<- "fail"
gradethis::grade_this({
  if (abs(.result - .solution) < 0.001) {
    exercise_results$prob_a_and_b <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::

Good work so far. Let's say you forgot to study for your chemistry quiz today. It is a four question multiple choice quiz with answer options $a)$ through $e)$ on each question. You decide that your best option is to guess randomly on each of the questions. Answer the following, using your spreadsheet to carry out any necessary computations.

:::{.callout-caution}
## Check Your Understanding: Guessing on a Quiz I

For a single question, what is the probability that you get that question correct?

```{ojs}
//| echo: false

viewof q6 = Inputs.radio(
  new Map([
    ["0", 1],
    ["0.10", 2],
    ["0.20", 3],
    ["0.25", 4],
    ["0.50", 5],
    ["0.80", 6],
    ["1", 7]
  ]),
  {value: JSON.parse(localStorage.getItem("q6_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q6_selected_t5d", JSON.stringify(q6));
  localStorage.setItem("q6_correct_t5d", "3");
  localStorage.setItem("q6_result_t5d", q6 === null ? "unattempted" : (q6 == 3 ? "correct" : "incorrect"));
}
ok_response(q6, "3");
```

:::

:::{.callout-caution}
## Check Your Understanding: Guessing on a Quiz II

What is the probability that you get every one of the questions correct? Use your spreadsheet to compute this, then select your answer.

```{ojs}
//| echo: false

viewof q7 = Inputs.radio(
  new Map([
    ["0", 1],
    ["0.0016", 2],
    ["0.01", 3],
    ["0.20", 4],
    ["0.4096", 5],
    ["0.5555", 6],
    ["0.998", 7],
    ["1", 8]
  ]),
  {value: JSON.parse(localStorage.getItem("q7_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q7_selected_t5d", JSON.stringify(q7));
  localStorage.setItem("q7_correct_t5d", "2");
  localStorage.setItem("q7_result_t5d", q7 === null ? "unattempted" : (q7 == 2 ? "correct" : "incorrect"));
}
ok_response(q7, "2");
```

:::{.hint exercise="q7-hint"}
:::{.callout-note collapse="false"}
## Hint 1

There are five answer options and only one is correct, so the probability of guessing correctly on one question is 1/5 = 0.20. The questions are independent of one another.

:::
:::

:::{.hint exercise="q7-hint"}
:::{.callout-note collapse="false"}
## Hint 2

Because the questions are independent, multiply the probability of getting each individual question correct together. In a spreadsheet cell, type `=0.20*0.20*0.20*0.20` and press Enter.

:::
:::

:::

:::{.callout-caution}
## Check Your Understanding: Guessing on a Quiz III

What is the probability that you get every one of the questions wrong? Use your spreadsheet to compute this, then select your answer.

```{ojs}
//| echo: false

viewof q8 = Inputs.radio(
  new Map([
    ["0", 1],
    ["0.0016", 2],
    ["0.01", 3],
    ["0.20", 4],
    ["0.4096", 5],
    ["0.5555", 6],
    ["0.998", 7],
    ["1", 8]
  ]),
  {value: JSON.parse(localStorage.getItem("q8_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q8_selected_t5d", JSON.stringify(q8));
  localStorage.setItem("q8_correct_t5d", "5");
  localStorage.setItem("q8_result_t5d", q8 === null ? "unattempted" : (q8 == 5 ? "correct" : "incorrect"));
}
ok_response(q8, "5");
```

:::{.hint exercise="q8-hint"}
:::{.callout-note collapse="false"}
## Hint 1

If the probability of getting a single question *correct* is 0.20, what is the probability of getting it *wrong*? These two outcomes cover every possibility, so they must add up to 1.

:::
:::

:::{.hint exercise="q8-hint"}
:::{.callout-note collapse="false"}
## Hint 2

The probability of getting a single question wrong is 1 - 0.20 = 0.80. Because the questions are independent, multiply this probability four times — once for each question. In your spreadsheet, type `=0.80*0.80*0.80*0.80` and press Enter.

:::
:::

:::

:::{.callout-caution}
## Check Your Understanding: Guessing on a Quiz IV

What is the probability that the first one you get wrong is question three? Use your spreadsheet to compute this, then select your answer.

```{ojs}
//| echo: false

viewof q9 = Inputs.radio(
  new Map([
    ["0", 1],
    ["0.015", 2],
    ["0.032", 3],
    ["0.0064", 4],
    ["0.20", 5],
    ["1.2", 6]
  ]),
  {value: JSON.parse(localStorage.getItem("q9_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q9_selected_t5d", JSON.stringify(q9));
  localStorage.setItem("q9_correct_t5d", "3");
  localStorage.setItem("q9_result_t5d", q9 === null ? "unattempted" : (q9 == 3 ? "correct" : "incorrect"));
}
ok_response(q9, "3");
```

:::{.hint exercise="q9-hint"}
:::{.callout-note collapse="false"}
## Hint 1

If the first incorrect answer is on question three, what must have happened on questions one and two? And what do we know about question four?

:::
:::

:::{.hint exercise="q9-hint"}
:::{.callout-note collapse="false"}
## Hint 2

Questions one and two must both have been answered correctly (probability 0.20 each), question three must have been answered incorrectly (probability 0.80), and question four's outcome doesn't matter — whether right or wrong, the event of interest has already occurred. The probability of either "right or wrong" on question four is 1.

:::
:::

:::{.hint exercise="q9-hint"}
:::{.callout-note collapse="false"}
## Hint 3

Multiply the four probabilities together. In your spreadsheet, type `=0.20*0.20*0.80*1` and press Enter.

:::
:::

:::

:::{.callout-caution}
## Check Your Understanding: Guessing on a Quiz V

What is the probability that you get exactly two questions right?

```{ojs}
//| echo: false

viewof q10 = Inputs.radio(
  new Map([
    ["0", 1],
    ["0.04", 2],
    ["0.08", 3],
    ["0.256", 4],
    ["0.20", 5],
    ["0.40", 6],
    ["0.64", 7],
    ["None of These", 8]
  ]),
  {value: JSON.parse(localStorage.getItem("q10_selected_t5d") ?? "null")}
);

{
  localStorage.setItem("q10_selected_t5d", JSON.stringify(q10));
  localStorage.setItem("q10_correct_t5d", "8");
  localStorage.setItem("q10_result_t5d", q10 === null ? "unattempted" : (q10 == 8 ? "correct" : "incorrect"));
}
ok_response(q10, "8");
```

:::

So in the last question, none of the choices were correct -- but why? There are lots of ways that we could get two of the questions right. We could get the first two right, the first and last right, the middle two right, and more! We need to account for all of these possibilities.

## Binomial Experiments and the Binomial Distribution

**Binomial Experiments:** A binomial experiment satisfies each of the following three criteria:

+ There are $n$ repeated trials.
+ Each trial has two possible outcomes (usually called *success* and *failure* for convenience).
+ The trials are independent of one another. That is, for each trial, the probability of success is $p$ (which remains constant).

**Binomial Distribution:** Let $X$ be the number of successes resulting from a binomial experiment with $n$ trials. We can compute the following probabilities using the `BINOM.DIST()` function in Excel or Google Sheets:

+ The probability of <u>exactly</u> $k$ successes is given by
<center> $\displaystyle{\mathbb{P}\left[X = k\right] = \binom{n}{k}\cdot p^k\left(1 - p\right)^{n-k} \approx \tt{BINOM.DIST(k,\ n,\ p,\ FALSE)}}$</center>
+ The probability of <u>at most</u> $k$ successes is given by
<center> $\displaystyle{\mathbb{P}\left[X \leq k\right] = \sum_{i=0}^{k}{\binom{n}{i}\cdot p^i\left(1 - p\right)^{n-i}} \approx \tt{BINOM.DIST(k,\ n,\ p,\ TRUE)}}$</center>

In the equations above, $\binom{n}{k} = \frac{n!}{k!\left(n-k\right)!}$ counts the number of ways to arrange the $k$ successes amongst the $n$ trials. The `BINOM.DIST()` function in your spreadsheet allows you to bypass the messy formulas -- but you'll still need to understand what the function does in order to use it correctly!

:::{.callout-tip}
## Tip: Using `BINOM.DIST()` in a Spreadsheet

The `BINOM.DIST()` function takes four arguments, in order:

1. **Number of successes** ($k$) — the outcome count you are interested in
2. **Number of trials** ($n$) — the total number of independent trials
3. **Probability of success** ($p$) — the probability of success on a single trial
4. **Cumulative** — use `FALSE` for *exactly* $k$ successes; use `TRUE` for *at most* $k$ successes

For example, `=BINOM.DIST(2, 7, 0.5, FALSE)` gives the probability of *exactly* 2 heads in 7 coin flips.
:::

:::{.callout-tip}
## Tip: Binomial Distribution

We need to use the binomial distribution to find probabilities associated with numbers of successful (or failing) outcomes in which *we do not know for certain the trials on which the successes (or failures) occur*.
:::

In your spreadsheet, use `BINOM.DIST()` to find the probability that you got exactly two of the four chemistry quiz questions from earlier correct. As a reminder, there were five answer options for each question and you were guessing randomly. Enter your result in the cell below.

```{webr}
#| exercise: chem-quiz-binom
# Enter your answer / result here

```

:::{.hint exercise="chem-quiz-binom"}
:::{.callout-note collapse="false"}
## Hint 1

In a blank spreadsheet cell, begin by typing `=BINOM.DIST(`. The four arguments you need to supply are: number of successes, number of trials, probability of success, and whether to use the cumulative form.

:::
:::

:::{.hint exercise="chem-quiz-binom"}
:::{.callout-note collapse="false"}
## Hint 2

We want *exactly* two correct, so the last argument should be `FALSE`. The first argument is `2` (the number of correct answers we're interested in). Fill in the other two.

:::
:::

:::{.hint exercise="chem-quiz-binom"}
:::{.callout-note collapse="false"}
## Hint 3

The quiz has four questions, so the number of trials is `4`. What is the probability of guessing correctly on a single question?

:::
:::

:::{.hint exercise="chem-quiz-binom"}
:::{.callout-note collapse="false"}
## Hint 4

Since each question has five possible choices (a–e), one of which is correct, the probability of guessing correctly on one question is 0.2. In your spreadsheet, type `=BINOM.DIST(2, 4, 0.2, FALSE)` and press Enter.

:::
:::

```{webr}
#| exercise: chem-quiz-binom
#| solution: true

dbinom(2, 4, 0.2)
```

```{webr}
#| exercise: chem-quiz-binom
#| check: true

exercise_results$chem_quiz_binom <<- "fail"
gradethis::grade_this({
  if (abs(.result - .solution) < 0.001) {
    exercise_results$chem_quiz_binom <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Good work. Now you'll get to try a few more problems! As you work through the next set of questions, you may want to check out [this example and solution](www/BinomialExample_Excel.pdf){target="_blank"}. Note that in that document, I mention that drawing a simple picture for each problem will help you decide which form of `BINOM.DIST()` to use and whether you might need to combine multiple calculations. This is a really important strategy that will help you develop an approach to solve each problem.

**Practice:** For each of the following, consider a scenario in which a random sample of `{ojs} n_students` students is asked (in private) whether they've failed to hand in at least one assignment this semester. We assume that about `{ojs} p_pct`% of students fail to hand in at least one assignment.

1. Given a single, randomly chosen student, what is the probability that the student will have failed to hand in at least one assignment this semester?

```{webr}
#| exercise: binom-assignments-ex1
#| input:
#|   - p_assign
# Enter your answer / result here

```

:::{.hint exercise="binom-assignments-ex1"}
:::{.callout-note collapse="false"}
## Hint 1

If we had 100 randomly chosen students, how many might you expect to have failed to hand in at least one assignment? How did you get that number? What mathematical operation did you perform?

:::
:::

:::{.hint exercise="binom-assignments-ex1"}
:::{.callout-note collapse="false"}
## Hint 2

The percentage given in the problem statement directly tells you the probability for a single randomly selected student. Convert `{ojs} p_pct`% to a decimal — no spreadsheet formula is needed for this one.

:::
:::

:::{.hint exercise="binom-assignments-ex1"}
:::{.callout-note collapse="false"}
## Hint 3

`{ojs} p_pct`% as a decimal is `{ojs} p_assign`. You can type this value directly into the answer cell.

:::
:::

```{webr}
#| exercise: binom-assignments-ex1
#| solution: true
#| input:
#|   - p_assign

p_assign
```

```{webr}
#| exercise: binom-assignments-ex1
#| check: true

exercise_results$binom_assignments_ex1 <<- "fail"
gradethis::grade_this({
  if (abs(.result - .solution) < 0.001) {
    exercise_results$binom_assignments_ex1 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

2. Find the probability that exactly 7 of the `{ojs} n_students` students have failed to hand in at least one assignment. Do the calculations in your spreadsheet, then enter your result below.

```{webr}
#| exercise: binom-assignments-ex2
#| input:
#|   - n_students
#|   - p_assign
# Enter your answer / result here

```

:::{.hint exercise="binom-assignments-ex2"}
:::{.callout-note collapse="false"}
## Hint 1

We don't know which specific students have failed to hand in an assignment — only the total count. This is a job for `BINOM.DIST()`. In a blank spreadsheet cell, start by typing `=BINOM.DIST(`.

:::
:::

:::{.hint exercise="binom-assignments-ex2"}
:::{.callout-note collapse="false"}
## Hint 2

We want the probability of *exactly* 7 students, so the last argument (cumulative) should be `FALSE`. The first argument is the number of successes we're interested in — that's `7` here.

:::
:::

:::{.hint exercise="binom-assignments-ex2"}
:::{.callout-note collapse="false"}
## Hint 3

The second argument is the total number of trials (students in the sample), which is `{ojs} n_students`. The third argument is the probability that any single student has failed to hand in an assignment, which is `{ojs} p_assign`.

:::
:::

:::{.hint exercise="binom-assignments-ex2"}
:::{.callout-note collapse="false"}
## Hint 4

In your spreadsheet, type `=BINOM.DIST(7, ___, ___, FALSE)`, where the second argument is `{ojs} n_students` and the third is `{ojs} p_assign`, then press Return. Enter the resulting value below.

:::
:::

```{webr}
#| exercise: binom-assignments-ex2
#| solution: true
#| input:
#|   - n_students
#|   - p_assign

dbinom(7, n_students, p_assign)
```

```{webr}
#| exercise: binom-assignments-ex2
#| check: true

exercise_results$binom_assignments_ex2 <<- "fail"
gradethis::grade_this({
  if (abs(.result - .solution) < 0.0001) {
    exercise_results$binom_assignments_ex2 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

3. Find the probability that at most `{ojs} thresh_atmost` of the `{ojs} n_students` students have failed to hand in at least one assignment. Do the calculations in your spreadsheet, then enter your result below.

```{webr}
#| exercise: binom-assignments-ex3
#| input:
#|   - n_students
#|   - p_assign
#|   - thresh_atmost
# Enter your answer / result here

```

:::{.hint exercise="binom-assignments-ex3"}
:::{.callout-note collapse="false"}
## Hint 1

We want the probability of a *range* of outcomes (0, 1, 2, … up to `{ojs} thresh_atmost` students), not just one specific count. The `BINOM.DIST()` function with its cumulative option can handle this directly. In a blank spreadsheet cell, start by typing `=BINOM.DIST(`.

:::
:::

:::{.hint exercise="binom-assignments-ex3"}
:::{.callout-note collapse="false"}
## Hint 2

"At most" means we want the cumulative probability up through `{ojs} thresh_atmost`. Set the last argument (cumulative) to `TRUE`. The first argument is the maximum number of successes you are willing to consider, which is `{ojs} thresh_atmost`.

:::
:::

:::{.hint exercise="binom-assignments-ex3"}
:::{.callout-note collapse="false"}
## Hint 3

The second argument is the number of students ( `{ojs} n_students` ) and the third is the probability of a single student having failed to hand in an assignment ( `{ojs} p_assign` ).

:::
:::

:::{.hint exercise="binom-assignments-ex3"}
:::{.callout-note collapse="false"}
## Hint 4 (Solved)

In your spreadsheet, type `=BINOM.DIST(___, ___, ___, TRUE)` where the first argument is filled by `{ojs} thresh_atmost`, the second is filled by `{ojs} n_students`, and the third is filled by `{ojs} p_assign`.

:::
:::

```{webr}
#| exercise: binom-assignments-ex3
#| solution: true
#| input:
#|   - n_students
#|   - p_assign
#|   - thresh_atmost

pbinom(thresh_atmost, n_students, p_assign)
```

```{webr}
#| exercise: binom-assignments-ex3
#| check: true

exercise_results$binom_assignments_ex3 <<- "fail"
gradethis::grade_this({
  if (abs(.result - .solution) < 0.0001) {
    exercise_results$binom_assignments_ex3 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

4. Find the probability that at least `{ojs} thresh_atleast` of the `{ojs} n_students` students have failed to hand in at least one assignment. Do the calculations in your spreadsheet, then enter your result below.

```{webr}
#| exercise: binom-assignments-ex4
#| input:
#|   - n_students
#|   - p_assign
#|   - thresh_atleast
# Enter your answer / result here

```

:::{.hint exercise="binom-assignments-ex4"}
:::{.callout-note collapse="false"}
## Hint 1

The `BINOM.DIST()` function with `TRUE` gives the probability of *at most* some number of successes. "At least `{ojs} thresh_atleast`" is not quite the same thing — think about how you might use what `BINOM.DIST()` gives you to get what you need.

:::
:::

:::{.hint exercise="binom-assignments-ex4"}
:::{.callout-note collapse="false"}
## Hint 2

Between 0 and `{ojs} n_students` students must fail to hand in an assignment — that covers every possibility, so the total probability is 1. If you can find the probability of everything *below* your threshold, you can subtract it from 1 to get what remains.

:::
:::

:::{.hint exercise="binom-assignments-ex4"}
:::{.callout-note collapse="false"}
## Hint 3

You want to subtract the probability of outcomes that are *too small* — that is, the probability of *fewer than* `{ojs} thresh_atleast` students. Be careful: "fewer than `{ojs} thresh_atleast`" means at most `{ojs} Number(thresh_atleast) - 1`. Try building a formula of the form `=1 - BINOM.DIST(___, ___, ___, TRUE)`.

:::
:::

:::{.hint exercise="binom-assignments-ex4"}
:::{.callout-note collapse="false"}
## Hint 4

In your spreadsheet formula, `=1 - BINOM.DIST(___, ___, ___, TRUE)`, the second blank should be filled in by `{ojs} n_students` and the third should be filled in by `{ojs} p_assign`. What should be placed into the first blank?

:::
:::

:::{.hint exercise="binom-assignments-ex4"}
:::{.callout-note collapse="false"}
## Hint 5

In your spreadsheet formula, `=1 - BINOM.DIST(___, ___, ___, TRUE)`, the second blank should be filled in by `{ojs} n_students` and the third should be filled in by `{ojs} p_assign`. Filling the first blank with `{ojs} thresh_atleast` doesn't give the correct result. Why not?

:::
:::

:::{.hint exercise="binom-assignments-ex4"}
:::{.callout-note collapse="false"}
## Hint 6

In your spreadsheet formula, `=1 - BINOM.DIST(___, ___, ___, TRUE)`, the second blank should be filled in by `{ojs} n_students` and the third should be filled in by `{ojs} p_assign`. Filling the first blank with `{ojs} thresh_atleast` doesn't give the correct result. Why not?

Using `{ojs} thresh_atleast` in that first blank removes the probability of at most `{ojs} thresh_atleast` students failing to hand in an assignment. This leaves us only with the probability that `{ojs} Number(thresh_atleast) + 1` or more students have failed to hand in an assignment.

:::
:::

:::{.hint exercise="binom-assignments-ex4"}
:::{.callout-note collapse="false"}
## Hint 6

In your spreadsheet formula, `=1 - BINOM.DIST(___, ___, ___, TRUE)`, the second blank should be filled in by `{ojs} n_students` and the third should be filled in by `{ojs} p_assign`. Fill the first blank with `{ojs} thresh_atleast - 1` to remove the probability associated with the events you aren't interested in.

:::
:::

```{webr}
#| exercise: binom-assignments-ex4
#| solution: true
#| input:
#|   - n_students
#|   - p_assign
#|   - thresh_atleast

1 - pbinom(thresh_atleast - 1, n_students, p_assign)
```

```{webr}
#| exercise: binom-assignments-ex4
#| check: true

exercise_results$binom_assignments_ex4 <<- "fail"
gradethis::grade_this({
  if (abs(.result - .solution) < 0.0001) {
    exercise_results$binom_assignments_ex4 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

5. Find the probability that between a minimum of `{ojs} thresh_between_lo` and a maximum of `{ojs} thresh_between_hi` out of the `{ojs} n_students` students have failed to hand in at least one assignment. Do the calculations in your spreadsheet, then enter your result below.

```{webr}
#| exercise: binom-assignments-ex5
#| input:
#|   - n_students
#|   - p_assign
#|   - thresh_between_lo
#|   - thresh_between_hi
# Enter your answer / result here

```

:::{.hint exercise="binom-assignments-ex5"}
:::{.callout-note collapse="false"}
## Hint 1

We want the probability of an outcome somewhere in the *middle* — not too few, not too many. A single `BINOM.DIST()` call with `TRUE` always gives everything from 0 up to some value, so you'll need to use two calls and combine them.

:::
:::

:::{.hint exercise="binom-assignments-ex5"}
:::{.callout-note collapse="false"}
## Hint 2

Start by computing a cumulative probability which is too big. Which probability is definitely larger than the answer we are seeking out: filling the first blank in `BINOM.DIST(___, ___, ___, TRUE)` with `{ojs} thresh_between_hi` or filling it with `{ojs} thresh_between_lo`? 

:::
:::

:::{.hint exercise="binom-assignments-ex5"}
:::{.callout-note collapse="false"}
## Hint 3

Filling the first blank in `BINOM.DIST(___, ___, ___, TRUE)` with `{ojs} thresh_between_hi` results in a probability which is definitely larger than the probability we are trying to find. That's because it includes all the outcomes we're interested in along with some others as well. 

We'll need to make the probability smaller.

:::
:::

:::{.hint exercise="binom-assignments-ex5"}
:::{.callout-note collapse="false"}
## Hint 4

Let's finish off that first call to `BINOM.DIST(___, ___, ___, TRUE)`. Fill the first blank with `{ojs} thresh_between_hi`, the second with `{ojs} n_students`, and the third with `{ojs} p_assign`.

:::
:::

:::{.hint exercise="binom-assignments-ex5"}
:::{.callout-note collapse="false"}
## Hint 5

Let's finish off that first call to `BINOM.DIST(___, ___, ___, TRUE)`. Fill the first blank with `{ojs} thresh_between_hi`, the second with `{ojs} n_students`, and the third with `{ojs} p_assign`.

We mentioned that this probability is too large, so we'll need to subtract some away.

:::
:::

:::{.hint exercise="binom-assignments-ex5"}
:::{.callout-note collapse="false"}
## Hint 6

Let's finish off that first call to `BINOM.DIST(___, ___, ___, TRUE)`. Fill the first blank with `{ojs} thresh_between_hi`, the second with `{ojs} n_students`, and the third with `{ojs} p_assign`.

We mentioned that this probability is too large, so we'll need to subtract some away. You'll use another call to `BINOM.DIST()`, so the formula you type into the cell will look like the following:

`BINOM.DIST(___, ___, ___, TRUE) - BINOM.DIST(___, ___, ___, TRUE)`

:::
:::

:::{.hint exercise="binom-assignments-ex5"}
:::{.callout-note collapse="false"}
## Hint 7

The formula you type into the cell will look like the following:

`BINOM.DIST(___, ___, ___, TRUE) - BINOM.DIST(___, ___, ___, TRUE)`

You'll fill in the second and third blanks in each call to `BINOM.DIST()` with `{ojs} n_students` and `{ojs} p_assign`. The first blank in the first call will be `{ojs} thresh_between_hi`. What should the first blank in the second call to `BINOM.DIST()` be?

:::
:::

:::{.hint exercise="binom-assignments-ex5"}
:::{.callout-note collapse="false"}
## Hint 8

The formula you type into the cell will look like the following:

`BINOM.DIST(___, ___, ___, TRUE) - BINOM.DIST(___, ___, ___, TRUE)`

You'll fill in the second and third blanks in each call to `BINOM.DIST()` with `{ojs} n_students` and `{ojs} p_assign`. The first blank in the first call will be `{ojs} thresh_between_hi`. What should the first blank in the second call to `BINOM.DIST()` be?

Filling that blank with `{ojs} thresh_between_lo` results in the wrong answer -- why?

:::
:::

:::{.hint exercise="binom-assignments-ex5"}
:::{.callout-note collapse="false"}
## Hint 9 (Solved)

The formula you type into the cell will look like the following:

`BINOM.DIST(___, ___, ___, TRUE) - BINOM.DIST(___, ___, ___, TRUE)`

You'll fill in the second and third blanks in each call to `BINOM.DIST()` with `{ojs} n_students` and `{ojs} p_assign`. The first blank in the first call will be `{ojs} thresh_between_hi`. What should the first blank in the second call to `BINOM.DIST()` be?

Filling that blank with `{ojs} thresh_between_lo` results in the wrong answer -- why? That removes the probability associated with one of the events we were interested in! Fill that blank with `{ojs} Number(thresh_between_lo) - 1` instead.

:::
:::

```{webr}
#| exercise: binom-assignments-ex5
#| solution: true
#| input:
#|   - n_students
#|   - p_assign
#|   - thresh_between_lo
#|   - thresh_between_hi

pbinom(thresh_between_hi, n_students, p_assign) - pbinom(thresh_between_lo - 1, n_students, p_assign)
```

```{webr}
#| exercise: binom-assignments-ex5
#| check: true

exercise_results$binom_assignments_ex5 <<- "fail"
gradethis::grade_this({
  if (abs(.result - .solution) < 0.0001) {
    exercise_results$binom_assignments_ex5 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-important}
## Don't Memorize Approaches

In several of the previous scenarios, we needed to think carefully about the correct first argument to pass to `BINOM.DIST()`, and whether to add or subtract a second call. Don't try to memorize when to subtract one from the threshold, when to leave it as-is, etc. The *language* of the problem is what matters, and there are lots of ways to express which outcomes we care about. If you insist on memorizing, you'll become frustrated quickly.

Instead of memorizing, take the time to *draw a picture* to help you. Examples of what these pictures might look like can be seen in [the example and solution document](www/BinomialExample_Excel.pdf){target="_blank"}, which I pointed you to earlier.
:::

6. The expected number of successes in a binomial experiment is sometimes denoted by $\mathbb{E}\left[X\right]$ and can be computed as $\mathbb{E}\left[X\right] = n\cdot p$, where $n$ denotes the number of trials run and $p$ denotes the probability of success on a single trial. Sometimes it is convenient to think of the expected number of successes as "the mean". Use your spreadsheet to compute the expected number of students (out of `{ojs} n_students`) who have failed to hand in at least one assignment, then enter your result below.

```{webr}
#| exercise: binom-assignments-ex6
#| input:
#|   - n_students
#|   - p_assign
# Enter your answer / result here

```

:::{.hint exercise="binom-assignments-ex6"}
:::{.callout-note collapse="false"}
## Hint 1

Use the formula $\mathbb{E}\left[X\right] = n\cdot p$ from the problem statement. In a blank spreadsheet cell, type `=` followed by the values of $n$ and $p$ multiplied together.

:::
:::

:::{.hint exercise="binom-assignments-ex6"}
:::{.callout-note collapse="false"}
## Hint 2

Here $n$ is the number of students in the sample (`{ojs} n_students`) and $p$ is the probability that any one student has failed to hand in an assignment (`{ojs} p_assign`).

:::
:::

:::{.hint exercise="binom-assignments-ex6"}
:::{.callout-note collapse="false"}
## Hint 3

In your spreadsheet, type = `{ojs} n_students` * `{ojs} p_assign` and press Return. Enter the resulting value below.

:::
:::

```{webr}
#| exercise: binom-assignments-ex6
#| solution: true
#| input:
#|   - n_students
#|   - p_assign

n_students * p_assign
```

```{webr}
#| exercise: binom-assignments-ex6
#| check: true

exercise_results$binom_assignments_ex6 <<- "fail"
gradethis::grade_this({
  if (abs(.result - .solution) < 0.01) {
    exercise_results$binom_assignments_ex6 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

7. The standard deviation in the number of successes for a binomial experiment can also be computed. The quantity $\displaystyle{s_X = \sqrt{n\cdot p\left(1 - p\right)}}$, where $n$ denotes the number of trials run and $p$ denotes the probability of success on a single trial, is the standard deviation in number of successes. Use your spreadsheet to compute the standard deviation in number of students who have failed to hand in at least one assignment from random samples of `{ojs} n_students` students, then enter your result below.

```{webr}
#| exercise: binom-assignments-ex7
#| input:
#|   - n_students
#|   - p_assign
# Enter your answer / result here

```

:::{.hint exercise="binom-assignments-ex7"}
:::{.callout-note collapse="false"}
## Hint 1

Use the formula $s_X = \sqrt{n\cdot p\left(1 - p\right)}$ from the problem statement. As a reminder, you can compute a square root in a spreadsheet using the `SQRT()` function.

:::
:::

:::{.hint exercise="binom-assignments-ex7"}
:::{.callout-note collapse="false"}
## Hint 2

For this problem, $n$ = `{ojs} n_students` and $p$ = `{ojs} p_assign`. Build your formula inside `SQRT()` by multiplying $n$, $p$, and $(1 - p)$ together.

:::
:::

:::{.hint exercise="binom-assignments-ex7"}
:::{.callout-note collapse="false"}
## Hint 3

In your spreadsheet, type =SQRT( `{ojs} n_students` * `{ojs} p_assign` * (1 - `{ojs} p_assign`)) and press Enter. Enter the resulting value below.

:::
:::

```{webr}
#| exercise: binom-assignments-ex7
#| solution: true
#| input:
#|   - n_students
#|   - p_assign

sqrt(n_students * p_assign * (1 - p_assign))
```

```{webr}
#| exercise: binom-assignments-ex7
#| check: true

exercise_results$binom_assignments_ex7 <<- "fail"
gradethis::grade_this({
  if (abs(.result - .solution) < 0.001) {
    exercise_results$binom_assignments_ex7 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Be sure to write down what questions you had as you worked through these problems and to have a teacher, colleague, or tutor help clarify things for you.

## Submit

If you are part of a course with an instructor who is grading your work on these activities, please copy and submit ***both*** of the hashes below using the method your instructor has requested.

:::{.callout-warning}
## Question Hash

The hash below encodes your responses to the multiple choice questions in this activity.

```{ojs}
//| echo: false

function buildQuestionResults() {
  return {
    notebook: "Topic 5: Discrete Probability Distributions",
    type: "questions",
    timestamp: new Date().toISOString(),
    questions: {
      q1_basic_prob_1: {
        selected: q1,
        correct_answer: "2",
        result: q1 === null ? "unattempted" : (q1 == 2 ? "correct" : "incorrect")
      },
      q2_basic_prob_2: {
        selected: q2,
        correct_answer: "4",
        result: q2 === null ? "unattempted" : (q2 == 4 ? "correct" : "incorrect")
      },
      q3_basic_prob_3: {
        selected: q3,
        correct_answer: "2",
        result: q3 === null ? "unattempted" : (q3 == 2 ? "correct" : "incorrect")
      },
      q4_independent_events_1: {
        selected: q4,
        correct_answer: "6",
        result: q4 === null ? "unattempted" : (q4 == 6 ? "correct" : "incorrect")
      },
      q5_independent_events_2: {
        selected: q5,
        correct_answer: "2",
        result: q5 === null ? "unattempted" : (q5 == 2 ? "correct" : "incorrect")
      },
      q6_quiz_1: {
        selected: q6,
        correct_answer: "3",
        result: q6 === null ? "unattempted" : (q6 == 3 ? "correct" : "incorrect")
      },
      q7_quiz_2: {
        selected: q7,
        correct_answer: "2",
        result: q7 === null ? "unattempted" : (q7 == 2 ? "correct" : "incorrect")
      },
      q8_quiz_3: {
        selected: q8,
        correct_answer: "5",
        result: q8 === null ? "unattempted" : (q8 == 5 ? "correct" : "incorrect")
      },
      q9_quiz_4: {
        selected: q9,
        correct_answer: "3",
        result: q9 === null ? "unattempted" : (q9 == 3 ? "correct" : "incorrect")
      },
      q10_quiz_5: {
        selected: q10,
        correct_answer: "8",
        result: q10 === null ? "unattempted" : (q10 == 8 ? "correct" : "incorrect")
      }
    }
  };
}

function toBase64(str) {
  return btoa(unescape(encodeURIComponent(str)));
}

question_hash = {
  q1; q2; q3; q4; q5; q6; q7; q8; q9; q10;
  return toBase64(JSON.stringify(buildQuestionResults()));
}
```

```{ojs}
//| echo: false

html`<div style="font-family: monospace; font-size: 0.85em; background: #f5f5f5; padding: 12px; border-radius: 6px; word-break: break-all; border: 1px solid #ddd; user-select: all; cursor: pointer;" onclick="navigator.clipboard.writeText(this.innerText)">
  ${question_hash}
</div>
<p style="margin-top: 8px; font-size: 0.9em; color: #555;">
  Click the box to copy to clipboard.
</p>`
```

:::

:::{.callout-warning}
## Exercise Hash

Click the button below to generate your exercise submission code. This hash encodes your work on the graded exercises in this activity.

You must have attempted the graded exercises before clicking — clicking generates a snapshot of your current results. If you have completed the activity over multiple sessions, please go back through and enter your answer in each graded exercise cell before generating the hash below, to ensure your most recent results are recorded.

```{webr}
#| exercise: generate-exercise-hash
#| caption: "Generate Exercise Code"

exercise_results$timestamp <- format(Sys.time(), "%Y-%m-%dT%H:%M:%SZ", tz = "UTC")
exercise_results$type <- "exercises"

hash <- base64encode(charToRaw(toJSON(exercise_results, auto_unbox = TRUE)))
cat(hash)
```

:::

## Summary

:::{.callout-tip}
## Main Takeaways

+ The **probability** of an event $A$ is a measure of its likelihood and is denoted $\mathbb{P}[A]$. Every probability must be between 0 and 1.
+ If $A$ and $B$ are **independent events**, then $\mathbb{P}[A \text{ and } B] = \mathbb{P}[A] \cdot \mathbb{P}[B]$.
+ A **binomial experiment** satisfies: (1) $n$ repeated trials, (2) each trial has two possible outcomes, and (3) trials are independent with constant probability of success $p$.
+ If $X$ counts successes in a binomial experiment with $n$ trials and success probability $p$, use `BINOM.DIST()` in your spreadsheet:
  + `=BINOM.DIST(k, n, p, FALSE)` — probability of **exactly** $k$ successes
  + `=BINOM.DIST(k, n, p, TRUE)` — probability of **at most** $k$ successes
  + Draw a picture to help you see how to combine these to handle "at least", "more than", or "between" scenarios. These two forms of `BINOM.DIST()` are sufficient to handle any binomial probability scenario — the challenge is identifying how to combine them.
+ The **expected number** of successes is $\mathbb{E}[X] = n \cdot p$.
+ The **standard deviation** of number of successes is $s_X = \sqrt{n \cdot p \cdot (1 - p)}$. Use `=SQRT(n * p * (1 - p))` in your spreadsheet.
:::

:::{.callout-tip}
## Looking Ahead

The next activity introduces the **normal distribution** — a continuous probability distribution that underpins much of classical statistical inference. Our focus will be on learning to compute probabilities and percentiles from this important distribution.
:::
