---
title: "Topic 16: Inference Practice (Part A)"
format: 
  live-html:
    include-in-header:
      text: |
        <style>
        .btn-exercise-solution { display: none !important; }
        .exercise-editor-body { max-height: 300px; overflow-y: auto; }
        </style>
    webr:
      packages:
        - dplyr
        - ggplot2
        - gradethis
      cell-options:
        persist: true
    live:
      show-solutions: true
---

{{< include ./_extensions/r-wasm/live/_knitr.qmd >}}
{{< include ./_extensions/r-wasm/live/_gradethis.qmd >}}

```{webr}
#| edit: false
#| output: false
#| define:
#|   - ok_response

library(htmltools)
ok_response <- function(response, n) {
  if (is.na(response)) div(HTML("You haven't answered yet."), style = "color: purple")
  else if (response == n) div(HTML("Correct ✓"), style = "color: green")
  else div(HTML("Incorrect ✗"), style = "color: red")
}
```

```{webr}
#| context: setup
#| echo: false
#| output: false

library(dplyr)
library(ggplot2)

theme_set(theme_bw(base_size = 12))
```

```{r setup, include=FALSE}
library(dplyr)
library(ggplot2)

theme_set(theme_bw(base_size = 12))
```

:::{.callout-note}
## About

No new content is introduced in this activity. This first of three practice notebooks contains four scaffolded inference problems involving probabilities, hypothesis tests, and confidence intervals for means and proportions. Use these problems to build your independence before moving on to unscaffolded homework problems.
:::

## Inference Practice

In this activity, you'll have an opportunity to practice a set of problems. Guiding questions are still provided for each problem to help you identify the necessary steps — think carefully about each step and how it relates to the larger inference process, since you'll need to navigate these steps independently on upcoming homework problems.

For convenience, the reference documents are linked here:

+ [*Standard Error Decision Tree*](https://agmath.github.io/SiteFiles/StdErrorDecisionTree.pdf){target="_blank"}
+ [*General Strategy for Confidence Intervals*](https://agmath.github.io/SiteFiles/CIgeneral.pdf){target="_blank"}
+ [*General Strategy for Hypothesis Tests*](https://agmath.github.io/SiteFiles/HTgeneral.pdf){target="_blank"}

For any problems involving hypothesis tests, assume $\alpha = 0.05$ unless otherwise stated.

:::{.callout-important}
## No Hints

There are no hints included in this activity. Use the reference documents linked above and your previous notebooks to guide you. These four problems include links to fully worked solutions for emergencies — but try to work through each problem on your own first. The next two practice activities won't have those supports.
:::

---

## Problem 1

The distribution of the number of eggs laid by a certain species of hen during their breeding period has a mean of 35 eggs and a standard deviation of 18.2. Suppose a group of researchers randomly samples 45 hens of this species, counts the number of eggs laid during their breeding period, and records a sample mean of 37 eggs. Find the probability of observing a sample of 45 hens whose mean number of eggs laid during the breeding period is at least 37. ([See worked solution](www/SampleInferenceProbSol.pdf){target="_blank"})

:::{.callout-caution}
## Problem 1, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

mutable ok_response = (response, n) => { return html`Loading...` };
viewof q1 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q1") ?? "null")}
);

localStorage.setItem("q1", JSON.stringify(q1));
ok_response(q1, "1");
```

:::

:::{.callout-caution}
## Problem 1, Part II

How many groups are sampled in this application?

```{ojs}
//| echo: false

viewof q2 = Inputs.radio(
  new Map([
    ["The sample was taken from a single population.", 1],
    ["Sample observations were collected from two populations.", 2],
    ["Sample observations were collected from more than two populations.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q2") ?? "null")}
);

localStorage.setItem("q2", JSON.stringify(q2));
ok_response(q2, "1");
```

:::

:::{.callout-caution}
## Problem 1, Part III

Do we know the population standard deviation for number of eggs laid?

```{ojs}
//| echo: false

viewof q3 = Inputs.radio(
  new Map([
    ["Yes, the population standard deviation is known.", 1],
    ["No, the population standard deviation is not known.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q3") ?? "null")}
);

localStorage.setItem("q3", JSON.stringify(q3));
ok_response(q3, "1");
```

:::

:::{.callout-caution}
## Problem 1, Part IV

How should the standard error of the sampling distribution be computed?

```{ojs}
//| echo: false

viewof q4 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q4") ?? "null")}
);

localStorage.setItem("q4", JSON.stringify(q4));
ok_response(q4, "1");
```

:::

:::{.callout-caution}
## Problem 1, Part V

Which distribution best describes the sampling distribution?

```{ojs}
//| echo: false

viewof q5 = Inputs.radio(
  new Map([
    ["A normal distribution.", 1],
    ["A t-distribution with n - 1 degrees of freedom.", 2],
    ["A t-distribution with min(n1-1, n2-1) degrees of freedom.", 3],
    ["A t-distribution with n_diff - 1 degrees of freedom.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q5") ?? "null")}
);

localStorage.setItem("q5", JSON.stringify(q5));
ok_response(q5, "1");
```

:::

Use the code block below to compute the probability of observing a random sample of 45 hens from this species which lays an average of at least 37 eggs.

```{webr}
#| exercise: prob1-5
______
```

```{webr}
#| exercise: prob1-5
#| solution: true

1 - pnorm(37, mean = 35, sd = 18.2/sqrt(45))
```

```{webr}
#| exercise: prob1-5
#| check: true

gradethis::grade_this({
  if (abs(.result - (1 - pnorm(37, 35, 18.2/sqrt(45)))) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

---

## Problem 2

A recent study looked at average ages of male and female inmates on Death Row. A random sample of 9 women and 14 men on Death Row resulted in the following:

| | Men | Women |
|:---|:---:|:---:|
| $\bar{x}$ | 39 | 44 |
| $s$ | 4.5 | 6.2 |

Conduct a hypothesis test to determine whether there is evidence to suggest that men on Death Row are, on average, younger than women on Death Row. ([See worked solution](www/SampleInferenceDeathRow.pdf){target="_blank"})

:::{.callout-caution}
## Problem 2, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

viewof q6 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q6") ?? "null")}
);

localStorage.setItem("q6", JSON.stringify(q6));
ok_response(q6, "3");
```

:::

:::{.callout-caution}
## Problem 2, Part II

What is the level of significance associated with this test?

```{ojs}
//| echo: false

viewof q7 = Inputs.radio(
  new Map([
    ["α = 0.01", 1],
    ["α = 0.05", 2],
    ["α = 0.10", 3],
    ["The p-value.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q7") ?? "null")}
);

localStorage.setItem("q7", JSON.stringify(q7));
ok_response(q7, "2");
```

:::

:::{.callout-caution}
## Problem 2, Part III

Does this hypothesis test involve testing a statement about a mean ($\mu$), a proportion ($p$), or something else?

```{ojs}
//| echo: false

viewof q8 = Inputs.radio(
  new Map([
    ["One or more means.", 1],
    ["One or more proportions.", 2],
    ["Something else altogether.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q8") ?? "null")}
);

localStorage.setItem("q8", JSON.stringify(q8));
ok_response(q8, "1");
```

:::

:::{.callout-caution}
## Problem 2, Part IV

How many groups are being compared in this test?

```{ojs}
//| echo: false

viewof q9 = Inputs.radio(
  new Map([
    ["The test involves only a single group.", 1],
    ["The test compares two groups.", 2],
    ["The test compares more than two groups.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q9") ?? "null")}
);

localStorage.setItem("q9", JSON.stringify(q9));
ok_response(q9, "2");
```

:::

:::{.callout-caution}
## Problem 2, Part V

Which of the following are the hypotheses associated with this test?

```{ojs}
//| echo: false

viewof q10 = Inputs.radio(
  new Map([
    ["H₀: μ_men − μ_women = 0; Hₐ: μ_men − μ_women ≠ 0", 1],
    ["H₀: μ_men − μ_women = 0; Hₐ: μ_men − μ_women > 0", 2],
    ["H₀: μ_men − μ_women = 0; Hₐ: μ_men − μ_women < 0", 3],
    ["H₀: μ_men ≠ μ_women; Hₐ: μ_men = μ_women", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q10") ?? "null")}
);

localStorage.setItem("q10", JSON.stringify(q10));
ok_response(q10, "3");
```

:::

:::{.callout-caution}
## Problem 2, Part VI

Do we know the population standard deviations ($\sigma$) for age in each group?

```{ojs}
//| echo: false

viewof q11 = Inputs.radio(
  new Map([
    ["Yes, the population standard deviations are known.", 1],
    ["No, the population standard deviations are not known.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q11") ?? "null")}
);

localStorage.setItem("q11", JSON.stringify(q11));
ok_response(q11, "2");
```

:::

:::{.callout-caution}
## Problem 2, Part VII

Are the observations in the two groups paired?

```{ojs}
//| echo: false

viewof q12 = Inputs.radio(
  new Map([
    ["No. There is no reason to suggest that the prisoners are paired.", 1],
    ["Yes. Every male prisoner can be naturally paired to a female prisoner.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q12") ?? "null")}
);

localStorage.setItem("q12", JSON.stringify(q12));
ok_response(q12, "1");
```

:::

:::{.callout-caution}
## Problem 2, Part VIII

Which standard error formula should be used?

```{ojs}
//| echo: false

viewof q13 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q13") ?? "null")}
);

localStorage.setItem("q13", JSON.stringify(q13));
ok_response(q13, "3");
```

:::

:::{.callout-caution}
## Problem 2, Part IX

Which distribution does the test statistic follow?

```{ojs}
//| echo: false

viewof q14 = Inputs.radio(
  new Map([
    ["The normal distribution.", 1],
    ["The t-distribution with df = n - 1.", 2],
    ["The t-distribution with df = min(n1 - 1, n2 - 1).", 3],
    ["The t-distribution with df = n_diff - 1.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q14") ?? "null")}
);

localStorage.setItem("q14", JSON.stringify(q14));
ok_response(q14, "3");
```

:::

Use the code block below to compute the point estimate.

```{webr}
#| exercise: prob2-9
______
```

```{webr}
#| exercise: prob2-9
#| solution: true

39 - 44
```

```{webr}
#| exercise: prob2-9
#| check: true

gradethis::grade_this({
  if (abs(.result) == 5) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the null value.

```{webr}
#| exercise: prob2-10
______
```

```{webr}
#| exercise: prob2-10
#| solution: true

0
```

```{webr}
#| exercise: prob2-10
#| check: true

gradethis::grade_this({
  if (.result == 0) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the standard error.

```{webr}
#| exercise: prob2-11
______
```

```{webr}
#| exercise: prob2-11
#| solution: true

sqrt((4.5^2 / 14) + (6.2^2 / 9))
```

```{webr}
#| exercise: prob2-11
#| check: true

gradethis::grade_this({
  if (abs(.result - sqrt((4.5^2/14) + (6.2^2/9))) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the test statistic.

```{webr}
#| exercise: prob2-12
______
```

```{webr}
#| exercise: prob2-12
#| solution: true

(39 - 44) / sqrt((4.5^2 / 14) + (6.2^2 / 9))
```

```{webr}
#| exercise: prob2-12
#| check: true

gradethis::grade_this({
  if (abs(abs(.result) - abs((39 - 44) / sqrt((4.5^2/14) + (6.2^2/9)))) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the $p$-value.

```{webr}
#| exercise: prob2-13
______
```

```{webr}
#| exercise: prob2-13
#| solution: true

ts <- (39 - 44) / sqrt((4.5^2 / 14) + (6.2^2 / 9))
pt(-abs(ts), df = 8)
```

```{webr}
#| exercise: prob2-13
#| check: true

gradethis::grade_this({
  if (abs(.result - .solution) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-caution}
## Problem 2, Part X

What is the result of the test?

```{ojs}
//| echo: false

viewof q15 = Inputs.radio(
  new Map([
    ["Since p ≥ α, we do not have enough evidence to reject the null hypothesis.", 1],
    ["Since p ≥ α, we accept the null hypothesis.", 2],
    ["Since p < α, we reject the null hypothesis and accept the alternative hypothesis.", 3],
    ["Since p < α, we fail to reject the null hypothesis.", 4],
    ["It is impossible to determine.", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q15") ?? "null")}
);

localStorage.setItem("q15", JSON.stringify(q15));
ok_response(q15, "3");
```

:::

:::{.callout-caution}
## Problem 2, Part XI

The result of the test means that:

```{ojs}
//| echo: false

viewof q16 = Inputs.radio(
  new Map([
    ["The sample data did not provide significant evidence to suggest that the average age of men on Death Row is younger than the average age of women on Death Row.", 1],
    ["The sample data proved that the average age of men on Death Row is not younger than the average age of women on Death Row.", 2],
    ["The sample data provided significant evidence to suggest that the average age of men on Death Row is younger than the average age of women on Death Row.", 3],
    ["The sample data provided significant evidence to suggest that the sample average age of men on Death Row is younger than the sample average age of women on Death Row.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q16") ?? "null")}
);

localStorage.setItem("q16", JSON.stringify(q16));
ok_response(q16, "3");
```

:::

---

## Problem 3

A recent random sample of 30 customers who switched to a car insurance company boasting that new customers save an average of \$534 per year resulted in an average savings of \$525 with a standard deviation of \$40. Find a 90% confidence interval for the true mean savings of customers switching to this company. ([See worked solution](www/SampleInferenceCarInsurance.pdf){target="_blank"})

:::{.callout-caution}
## Problem 3, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

viewof q17 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q17") ?? "null")}
);

localStorage.setItem("q17", JSON.stringify(q17));
ok_response(q17, "2");
```

:::

:::{.callout-caution}
## Problem 3, Part II

What is the desired level of confidence?

```{ojs}
//| echo: false

viewof q18 = Inputs.radio(
  new Map([
    ["90%", 1],
    ["95%", 2],
    ["98%", 3],
    ["99%", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q18") ?? "null")}
);

localStorage.setItem("q18", JSON.stringify(q18));
ok_response(q18, "1");
```

:::

:::{.callout-caution}
## Problem 3, Part III

Is your confidence interval being built to capture a mean ($\mu$), a proportion ($p$), or something else?

```{ojs}
//| echo: false

viewof q19 = Inputs.radio(
  new Map([
    ["A mean or difference in means.", 1],
    ["A proportion or difference in proportions.", 2],
    ["Something else altogether.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q19") ?? "null")}
);

localStorage.setItem("q19", JSON.stringify(q19));
ok_response(q19, "1");
```

:::

:::{.callout-caution}
## Problem 3, Part IV

Does the population parameter belong to a single group or is it a comparison of multiple groups?

```{ojs}
//| echo: false

viewof q20 = Inputs.radio(
  new Map([
    ["The population parameter belongs to a single group.", 1],
    ["The population parameter involves a comparison of two groups.", 2],
    ["The population parameter involves a comparison of more than two groups.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q20") ?? "null")}
);

localStorage.setItem("q20", JSON.stringify(q20));
ok_response(q20, "1");
```

:::

:::{.callout-caution}
## Problem 3, Part V

Do we know the population standard deviation ($\sigma$) for dollars saved?

```{ojs}
//| echo: false

viewof q21 = Inputs.radio(
  new Map([
    ["Yes, the population standard deviation is known.", 1],
    ["No, the population standard deviation is not known.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q21") ?? "null")}
);

localStorage.setItem("q21", JSON.stringify(q21));
ok_response(q21, "2");
```

:::

:::{.callout-caution}
## Problem 3, Part VI

Which standard error formula should be used?

```{ojs}
//| echo: false

viewof q22 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q22") ?? "null")}
);

localStorage.setItem("q22", JSON.stringify(q22));
ok_response(q22, "2");
```

:::

:::{.callout-caution}
## Problem 3, Part VII

Which distribution should be used to identify the critical value?

```{ojs}
//| echo: false

viewof q23 = Inputs.radio(
  new Map([
    ["The normal distribution.", 1],
    ["The t-distribution with df = n - 1.", 2],
    ["The t-distribution with df = min(n1-1, n2-1).", 3],
    ["The t-distribution with df = n_diff - 1.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q23") ?? "null")}
);

localStorage.setItem("q23", JSON.stringify(q23));
ok_response(q23, "2");
```

:::

Use the code block below to compute the critical value.

```{webr}
#| exercise: prob3-7
______
```

```{webr}
#| exercise: prob3-7
#| solution: true

qt(0.95, df = 29)
```

```{webr}
#| exercise: prob3-7
#| check: true

gradethis::grade_this({
  if (abs(abs(.result) - qt(0.95, df = 29)) < 0.005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the point estimate.

```{webr}
#| exercise: prob3-8
______
```

```{webr}
#| exercise: prob3-8
#| solution: true

525
```

```{webr}
#| exercise: prob3-8
#| check: true

gradethis::grade_this({
  if (.result == 525) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the standard error.

```{webr}
#| exercise: prob3-9
______
```

```{webr}
#| exercise: prob3-9
#| solution: true

40 / sqrt(30)
```

```{webr}
#| exercise: prob3-9
#| check: true

gradethis::grade_this({
  if (abs(.result - (40 / sqrt(30))) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the lower bound of the confidence interval.

```{webr}
#| exercise: prob3-10
______
```

```{webr}
#| exercise: prob3-10
#| solution: true

525 - qt(0.95, df = 29) * (40 / sqrt(30))
```

```{webr}
#| exercise: prob3-10
#| check: true

gradethis::grade_this({
  if (abs(.result - (525 - qt(0.95, df = 29) * 40 / sqrt(30))) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the upper bound of the confidence interval.

```{webr}
#| exercise: prob3-11
______
```

```{webr}
#| exercise: prob3-11
#| solution: true

525 + qt(0.95, df = 29) * (40 / sqrt(30))
```

```{webr}
#| exercise: prob3-11
#| check: true

gradethis::grade_this({
  if (abs(.result - (525 + qt(0.95, df = 29) * 40 / sqrt(30))) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-caution}
## Problem 3, Part VIII

The correct interpretation of this confidence interval is:

```{ojs}
//| echo: false

viewof q24 = Inputs.radio(
  new Map([
    ["We are 90% confident that the true average savings customers should expect when switching to this provider is between the lower and upper bounds calculated.", 1],
    ["90% of the savings amounts customers realize when switching to this provider will fall between the lower and upper bounds calculated.", 2],
    ["We are 90% confident that the average savings customers in this sample experienced when switching to this provider fall between the lower and upper bounds calculated.", 3],
    ["There is a 90% chance a customer switching to this provider will realize savings between the lower and upper bounds calculated.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q24") ?? "null")}
);

localStorage.setItem("q24", JSON.stringify(q24));
ok_response(q24, "1");
```

:::

---

## Problem 4

In January 2011, The Marist Poll published a report stating that 66% of adults nationally think licensed drivers should be required to retake their road test once they reach 65 years of age. In that same year, 200 random citizens of New Hampshire were asked whether they were in favor of an additional road test for drivers at 65 years of age — 118 responded that they were in favor. Construct a 95% confidence interval for the true proportion of New Hampshirites who are in favor of this proposal. ([See worked solution](www/SampleInferenceMaristPoll.pdf){target="_blank"})

:::{.callout-caution}
## Problem 4, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

viewof q25 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q25") ?? "null")}
);

localStorage.setItem("q25", JSON.stringify(q25));
ok_response(q25, "2");
```

:::

:::{.callout-caution}
## Problem 4, Part II

What is the desired level of confidence?

```{ojs}
//| echo: false

viewof q26 = Inputs.radio(
  new Map([
    ["90%", 1],
    ["95%", 2],
    ["98%", 3],
    ["99%", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q26") ?? "null")}
);

localStorage.setItem("q26", JSON.stringify(q26));
ok_response(q26, "2");
```

:::

:::{.callout-caution}
## Problem 4, Part III

Is your confidence interval being built to capture a mean ($\mu$), a proportion ($p$), or something else?

```{ojs}
//| echo: false

viewof q27 = Inputs.radio(
  new Map([
    ["A mean or difference in means.", 1],
    ["A proportion or difference in proportions.", 2],
    ["Something else altogether.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q27") ?? "null")}
);

localStorage.setItem("q27", JSON.stringify(q27));
ok_response(q27, "2");
```

:::

:::{.callout-caution}
## Problem 4, Part IV

Does the population parameter belong to a single group or is it a comparison of multiple groups?

```{ojs}
//| echo: false

viewof q28 = Inputs.radio(
  new Map([
    ["The population parameter belongs to a single group.", 1],
    ["The population parameter involves a comparison of two groups.", 2],
    ["The population parameter involves a comparison of more than two groups.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q28") ?? "null")}
);

localStorage.setItem("q28", JSON.stringify(q28));
ok_response(q28, "1");
```

:::

:::{.callout-caution}
## Problem 4, Part V

Which standard error formula should be used?

```{ojs}
//| echo: false

viewof q29 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q29") ?? "null")}
);

localStorage.setItem("q29", JSON.stringify(q29));
ok_response(q29, "4");
```

:::

:::{.callout-caution}
## Problem 4, Part VI

Which distribution should be used to identify the critical value?

```{ojs}
//| echo: false

viewof q30 = Inputs.radio(
  new Map([
    ["The normal distribution.", 1],
    ["The t-distribution with df = n - 1.", 2],
    ["The t-distribution with df = min(n1-1, n2-1).", 3],
    ["The t-distribution with df = n_diff - 1.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q30") ?? "null")}
);

localStorage.setItem("q30", JSON.stringify(q30));
ok_response(q30, "1");
```

:::

Use the code block below to compute the critical value.

```{webr}
#| exercise: prob4-6
______
```

```{webr}
#| exercise: prob4-6
#| solution: true

qnorm(0.975)
```

```{webr}
#| exercise: prob4-6
#| check: true

gradethis::grade_this({
  if (abs(abs(.result) - qnorm(0.975)) < 0.005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the point estimate.

```{webr}
#| exercise: prob4-7
______
```

```{webr}
#| exercise: prob4-7
#| solution: true

118 / 200
```

```{webr}
#| exercise: prob4-7
#| check: true

gradethis::grade_this({
  if (abs(.result - 118/200) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the standard error.

```{webr}
#| exercise: prob4-8
______
```

```{webr}
#| exercise: prob4-8
#| solution: true

sqrt((118/200) * (1 - 118/200) / 200)
```

```{webr}
#| exercise: prob4-8
#| check: true

gradethis::grade_this({
  p_hat <- 118/200
  if (abs(.result - sqrt(p_hat * (1 - p_hat) / 200)) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the lower bound of the confidence interval.

```{webr}
#| exercise: prob4-9
______
```

```{webr}
#| exercise: prob4-9
#| solution: true

p_hat <- 118/200
p_hat - qnorm(0.975) * sqrt(p_hat * (1 - p_hat) / 200)
```

```{webr}
#| exercise: prob4-9
#| check: true

gradethis::grade_this({
  p_hat <- 118/200
  lower <- p_hat - qnorm(0.975) * sqrt(p_hat * (1 - p_hat) / 200)
  if (abs(.result - lower) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the upper bound of the confidence interval.

```{webr}
#| exercise: prob4-10
______
```

```{webr}
#| exercise: prob4-10
#| solution: true

p_hat <- 118/200
p_hat + qnorm(0.975) * sqrt(p_hat * (1 - p_hat) / 200)
```

```{webr}
#| exercise: prob4-10
#| check: true

gradethis::grade_this({
  p_hat <- 118/200
  upper <- p_hat + qnorm(0.975) * sqrt(p_hat * (1 - p_hat) / 200)
  if (abs(.result - upper) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-caution}
## Problem 4, Part VII

The correct interpretation of this confidence interval is:

```{ojs}
//| echo: false

viewof q31 = Inputs.radio(
  new Map([
    ["We are 95% confident that the sample proportion of New Hampshirites favoring the policy falls between the lower and upper bounds calculated.", 1],
    ["95% of New Hampshirites favor the policy.", 2],
    ["We are 95% confident that the true proportion of New Hampshirites favoring the policy is between the lower and upper bounds calculated.", 3],
    ["There is a 95% chance the true proportion of New Hampshirites favoring the policy falls between the lower and upper bounds calculated.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q31") ?? "null")}
);

localStorage.setItem("q31", JSON.stringify(q31));
ok_response(q31, "3");
```

:::

## Submit

:::{.callout-warning}

Grading function is not currently available for these web-based notebooks. If you would like to utilize grading functionality, please use either the [Posit Cloud pre-built instance](https://posit.cloud/content/6328402){target="_blank"} or [install the `{learnr}` notebooks locally](https://agmath.github.io/IntroductoryStatistics/AccessingInteractiveNotes.html){target="_blank"}.

:::

## Summary

:::{.callout-tip}
## Main Takeaways

These four problems covered the range of inference tasks you'll encounter most frequently. A few reminders worth carrying forward:

+ **Always identify the task first** — probability, confidence interval, hypothesis test, or sample size. The task determines everything that follows.
+ **The decision tree is your guide** for identifying the correct standard error formula and distribution. The key questions: means or proportions? one group or two? population standard deviation known or unknown?
+ **When $\sigma$ is unknown**, use $s$ in the standard error formula and the $t$-distribution for critical values and $p$-values. When $\sigma$ is known, use the normal distribution.
+ **Interpretation matters.** A confidence interval captures a population parameter with a stated level of confidence. It is not a probability statement about the parameter, and not a statement about individual observations.
:::

:::{.callout-tip}
## Looking Ahead

Part B continues with four more mixed inference problems. The scaffolding remains, but the worked solution links are removed — you're increasingly on your own. Take stock of what is difficult for you and be sure to ask a teacher, mentor, or friend when you need help. Keep the decision tree and general strategy documents close at hand. You'll find them to be helpful.
:::
