---
title: "Topic 16: Inference Practice (Part A, Spreadsheets)"
format: 
  live-html:
    include-in-header:
      text: |
        <style>
        .btn-exercise-solution { display: none !important; }
        .exercise-editor-body { max-height: 300px; overflow-y: auto; }
        .exercise-editor-btn-run-code .btn-label-exercise-editor {
          font-size: 0;
        }
        .exercise-editor-btn-run-code .btn-label-exercise-editor::after {
          content: "Submit Answer";
          font-size: 1rem;
        }
        </style>
    webr:
      packages:
        - dplyr
        - ggplot2
        - gradethis
        - jsonlite
        - base64enc
      cell-options:
        persist: true
    live:
      show-solutions: true
---

{{< include ./_extensions/r-wasm/live/_knitr.qmd >}}
{{< include ./_extensions/r-wasm/live/_gradethis.qmd >}}

```{webr}
#| edit: false
#| output: false
#| define:
#|   - ok_response

library(htmltools)
ok_response <- function(response, n) {
  if (is.na(response)) div(HTML("You haven't answered yet."), style = "color: purple")
  else if (response == n) div(HTML("Correct ✓"), style = "color: green")
  else div(HTML("Incorrect ✗"), style = "color: red")
}
```

```{webr}
#| context: setup
#| echo: false
#| output: false

library(dplyr)
library(ggplot2)
library(jsonlite)
library(base64enc)

theme_set(theme_bw(base_size = 12))

exercise_results <- list(
  notebook    = "Topic 16: Inference Practice (Part A, Spreadsheets)",
  prob1_5     = "unattempted",
  prob2_9     = "unattempted",
  prob2_10    = "unattempted",
  prob2_11    = "unattempted",
  prob2_12    = "unattempted",
  prob2_13    = "unattempted",
  prob3_7     = "unattempted",
  prob3_8     = "unattempted",
  prob3_9     = "unattempted",
  prob3_10    = "unattempted",
  prob3_11    = "unattempted",
  prob4_6     = "unattempted",
  prob4_7     = "unattempted",
  prob4_8     = "unattempted",
  prob4_9     = "unattempted",
  prob4_10    = "unattempted"
)
```

```{r setup, include=FALSE}
library(dplyr)
library(ggplot2)

theme_set(theme_bw(base_size = 12))
```

:::{.callout-note}
## About

No new content is introduced in this activity. This first of three practice notebooks contains four scaffolded inference problems involving probabilities, hypothesis tests, and confidence intervals for means and proportions. Use these problems to build your independence before moving on to unscaffolded homework problems.

```{ojs}
//| echo: false

{
  const btn = html`<button style="
    background: #436b95;
    color: white;
    border: none;
    padding: 8px 16px;
    border-radius: 4px;
    cursor: pointer;
    font-size: 0.9em;
  ">🔄 Reset My Responses</button>`;

  btn.onclick = () => {
    if (confirm('Reset all responses for this activity? This cannot be undone.')) {
      Object.keys(localStorage)
        .filter(k => k.endsWith('_selected_t16ae') || k.endsWith('_correct_t16ae') || k.endsWith('_result_t16ae'))
        .forEach(k => localStorage.removeItem(k));
      location.reload();
    }
  };

  return btn;
}
```

*Note.* The button above resets multiple choice and checkbox questions. Currently, resetting answer entry cells must be done manually via hitting the *Start Over* button on each individual interactive cell.

:::

:::{.callout-note}
## Using a Spreadsheet Program

All calculations in this activity should be done in your spreadsheet. Open a blank sheet now and keep it alongside this activity — label each calculation clearly so you can follow your own work.

+ If you are a Google Sheets user, [open a blank Google Sheet](https://sheets.new){target="_blank"}
+ If you are an Excel user, open an empty workbook on your computer or [create a blank workbook in Excel online](https://office.live.com/start/Excel.aspx)
:::

## Inference Practice

In this activity, you'll have an opportunity to practice a set of problems. Guiding questions are still provided for each problem to help you identify the necessary steps — think carefully about each step and how it relates to the larger inference process, since you'll need to navigate these steps independently on upcoming homework problems.

For convenience, the reference documents are linked here:

+ [*Standard Error Decision Tree*](https://agmath.github.io/SiteFiles/StdErrorDecisionTree.pdf){target="_blank"}
+ [*General Strategy for Confidence Intervals*](https://agmath.github.io/SiteFiles/CIgeneral.pdf){target="_blank"}
+ [*General Strategy for Hypothesis Tests*](https://agmath.github.io/SiteFiles/HTgeneral.pdf){target="_blank"}

For any problems involving hypothesis tests, assume $\alpha = 0.05$ unless otherwise stated.

:::{.callout-important}
## No Hints

There are no hints included in this activity. Use the reference documents linked above and your previous notebooks to guide you. These four problems include links to fully worked solutions for emergencies — but try to work through each problem on your own first. The next two practice activities won't have those supports.
:::


---

## Problem 1

The distribution of the number of eggs laid by a certain species of hen during their breeding period has a mean of 35 eggs and a standard deviation of 18.2. Suppose a group of researchers randomly samples 45 hens of this species, counts the number of eggs laid during their breeding period, and records a sample mean of 37 eggs. Find the probability of observing a sample of 45 hens whose mean number of eggs laid during the breeding period is at least 37. ([See worked solution](www/SampleInferenceProbSol_Spreadsheets.pdf){target="_blank"})

:::{.callout-caution}
## Problem 1, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

mutable ok_response = (response, n) => { return html`Loading...` };
viewof q1 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q1_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q1_selected_t16ae", JSON.stringify(q1));
  localStorage.setItem("q1_correct_t16ae", "1");
  localStorage.setItem("q1_result_t16ae", q1 === null ? "unattempted" : (q1 == 1 ? "correct" : "incorrect"));
}
ok_response(q1, "1");
```

:::

:::{.callout-caution}
## Problem 1, Part II

How many groups are sampled in this application?

```{ojs}
//| echo: false

viewof q2 = Inputs.radio(
  new Map([
    ["The sample was taken from a single population.", 1],
    ["Sample observations were collected from two populations.", 2],
    ["Sample observations were collected from more than two populations.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q2_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q2_selected_t16ae", JSON.stringify(q2));
  localStorage.setItem("q2_correct_t16ae", "1");
  localStorage.setItem("q2_result_t16ae", q2 === null ? "unattempted" : (q2 == 1 ? "correct" : "incorrect"));
}
ok_response(q2, "1");
```

:::

:::{.callout-caution}
## Problem 1, Part III

Do we know the population standard deviation for number of eggs laid?

```{ojs}
//| echo: false

viewof q3 = Inputs.radio(
  new Map([
    ["Yes, the population standard deviation is known.", 1],
    ["No, the population standard deviation is not known.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q3_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q3_selected_t16ae", JSON.stringify(q3));
  localStorage.setItem("q3_correct_t16ae", "1");
  localStorage.setItem("q3_result_t16ae", q3 === null ? "unattempted" : (q3 == 1 ? "correct" : "incorrect"));
}
ok_response(q3, "1");
```

:::

:::{.callout-caution}
## Problem 1, Part IV

How should the standard error of the sampling distribution be computed?

```{ojs}
//| echo: false

viewof q4 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q4_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q4_selected_t16ae", JSON.stringify(q4));
  localStorage.setItem("q4_correct_t16ae", "1");
  localStorage.setItem("q4_result_t16ae", q4 === null ? "unattempted" : (q4 == 1 ? "correct" : "incorrect"));
}
ok_response(q4, "1");
```

:::

:::{.callout-caution}
## Problem 1, Part V

Which distribution best describes the sampling distribution?

```{ojs}
//| echo: false

viewof q5 = Inputs.radio(
  new Map([
    ["A normal distribution.", 1],
    ["A t-distribution with n - 1 degrees of freedom.", 2],
    ["A t-distribution with min(n1-1, n2-1) degrees of freedom.", 3],
    ["A t-distribution with n_diff - 1 degrees of freedom.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q5_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q5_selected_t16ae", JSON.stringify(q5));
  localStorage.setItem("q5_correct_t16ae", "1");
  localStorage.setItem("q5_result_t16ae", q5 === null ? "unattempted" : (q5 == 1 ? "correct" : "incorrect"));
}
ok_response(q5, "1");
```

:::

Use your spreadsheet to compute the probability of observing a random sample of 45 hens from this species which lays an average of at least 37 eggs. Enter your result below.

```{webr}
#| exercise: prob1-5
# Enter your answer / result here

```

```{webr}
#| exercise: prob1-5
#| solution: true

1 - pnorm(37, mean = 35, sd = 18.2/sqrt(45))
```

```{webr}
#| exercise: prob1-5
#| check: true

exercise_results$prob1_5 <<- "fail"
gradethis::grade_this({
  if (abs(.result - (1 - pnorm(37, 35, 18.2/sqrt(45)))) < 0.0005) {
    exercise_results$prob1_5 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```


---

## Problem 2

A recent study looked at average ages of male and female inmates on Death Row. A random sample of 9 women and 14 men on Death Row resulted in the following:

| | Men | Women |
|:---|:---:|:---:|
| $\bar{x}$ | 39 | 44 |
| $s$ | 4.5 | 6.2 |

Conduct a hypothesis test to determine whether there is evidence to suggest that men on Death Row are, on average, younger than women on Death Row. ([See worked solution](www/SampleInferenceDeathRow_Spreadsheets.pdf){target="_blank"})

:::{.callout-caution}
## Problem 2, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

viewof q6 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q6_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q6_selected_t16ae", JSON.stringify(q6));
  localStorage.setItem("q6_correct_t16ae", "3");
  localStorage.setItem("q6_result_t16ae", q6 === null ? "unattempted" : (q6 == 3 ? "correct" : "incorrect"));
}
ok_response(q6, "3");
```

:::

:::{.callout-caution}
## Problem 2, Part II

What is the level of significance associated with this test?

```{ojs}
//| echo: false

viewof q7 = Inputs.radio(
  new Map([
    ["α = 0.01", 1],
    ["α = 0.05", 2],
    ["α = 0.10", 3],
    ["The p-value.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q7_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q7_selected_t16ae", JSON.stringify(q7));
  localStorage.setItem("q7_correct_t16ae", "2");
  localStorage.setItem("q7_result_t16ae", q7 === null ? "unattempted" : (q7 == 2 ? "correct" : "incorrect"));
}
ok_response(q7, "2");
```

:::

:::{.callout-caution}
## Problem 2, Part III

Does this hypothesis test involve testing a statement about a mean ($\mu$), a proportion ($p$), or something else?

```{ojs}
//| echo: false

viewof q8 = Inputs.radio(
  new Map([
    ["One or more means.", 1],
    ["One or more proportions.", 2],
    ["Something else altogether.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q8_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q8_selected_t16ae", JSON.stringify(q8));
  localStorage.setItem("q8_correct_t16ae", "1");
  localStorage.setItem("q8_result_t16ae", q8 === null ? "unattempted" : (q8 == 1 ? "correct" : "incorrect"));
}
ok_response(q8, "1");
```

:::

:::{.callout-caution}
## Problem 2, Part IV

How many groups are being compared in this test?

```{ojs}
//| echo: false

viewof q9 = Inputs.radio(
  new Map([
    ["The test involves only a single group.", 1],
    ["The test compares two groups.", 2],
    ["The test compares more than two groups.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q9_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q9_selected_t16ae", JSON.stringify(q9));
  localStorage.setItem("q9_correct_t16ae", "2");
  localStorage.setItem("q9_result_t16ae", q9 === null ? "unattempted" : (q9 == 2 ? "correct" : "incorrect"));
}
ok_response(q9, "2");
```

:::

:::{.callout-caution}
## Problem 2, Part V

Which of the following are the hypotheses associated with this test?

```{ojs}
//| echo: false

viewof q10 = Inputs.radio(
  new Map([
    ["H₀: μ_men − μ_women = 0; Hₐ: μ_men − μ_women ≠ 0", 1],
    ["H₀: μ_men − μ_women = 0; Hₐ: μ_men − μ_women > 0", 2],
    ["H₀: μ_men − μ_women = 0; Hₐ: μ_men − μ_women < 0", 3],
    ["H₀: μ_men ≠ μ_women; Hₐ: μ_men = μ_women", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q10_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q10_selected_t16ae", JSON.stringify(q10));
  localStorage.setItem("q10_correct_t16ae", "3");
  localStorage.setItem("q10_result_t16ae", q10 === null ? "unattempted" : (q10 == 3 ? "correct" : "incorrect"));
}
ok_response(q10, "3");
```

:::

:::{.callout-caution}
## Problem 2, Part VI

Do we know the population standard deviations ($\sigma$) for age in each group?

```{ojs}
//| echo: false

viewof q11 = Inputs.radio(
  new Map([
    ["Yes, the population standard deviations are known.", 1],
    ["No, the population standard deviations are not known.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q11_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q11_selected_t16ae", JSON.stringify(q11));
  localStorage.setItem("q11_correct_t16ae", "2");
  localStorage.setItem("q11_result_t16ae", q11 === null ? "unattempted" : (q11 == 2 ? "correct" : "incorrect"));
}
ok_response(q11, "2");
```

:::

:::{.callout-caution}
## Problem 2, Part VII

Are the observations in the two groups paired?

```{ojs}
//| echo: false

viewof q12 = Inputs.radio(
  new Map([
    ["No. There is no reason to suggest that the prisoners are paired.", 1],
    ["Yes. Every male prisoner can be naturally paired to a female prisoner.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q12_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q12_selected_t16ae", JSON.stringify(q12));
  localStorage.setItem("q12_correct_t16ae", "1");
  localStorage.setItem("q12_result_t16ae", q12 === null ? "unattempted" : (q12 == 1 ? "correct" : "incorrect"));
}
ok_response(q12, "1");
```

:::

:::{.callout-caution}
## Problem 2, Part VIII

Which standard error formula should be used?

```{ojs}
//| echo: false

viewof q13 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q13_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q13_selected_t16ae", JSON.stringify(q13));
  localStorage.setItem("q13_correct_t16ae", "3");
  localStorage.setItem("q13_result_t16ae", q13 === null ? "unattempted" : (q13 == 3 ? "correct" : "incorrect"));
}
ok_response(q13, "3");
```

:::

:::{.callout-caution}
## Problem 2, Part IX

Which distribution does the test statistic follow?

```{ojs}
//| echo: false

viewof q14 = Inputs.radio(
  new Map([
    ["The normal distribution.", 1],
    ["The t-distribution with df = n - 1.", 2],
    ["The t-distribution with df = min(n1 - 1, n2 - 1).", 3],
    ["The t-distribution with df = n_diff - 1.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q14_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q14_selected_t16ae", JSON.stringify(q14));
  localStorage.setItem("q14_correct_t16ae", "3");
  localStorage.setItem("q14_result_t16ae", q14 === null ? "unattempted" : (q14 == 3 ? "correct" : "incorrect"));
}
ok_response(q14, "3");
```

:::

Use your spreadsheet to compute the point estimate and enter your result below.

```{webr}
#| exercise: prob2-9
# Enter your answer / result here

```

```{webr}
#| exercise: prob2-9
#| solution: true

39 - 44
```

```{webr}
#| exercise: prob2-9
#| check: true

exercise_results$prob2_9 <<- "fail"
gradethis::grade_this({
  if (abs(.result) == 5) {
    exercise_results$prob2_9 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the null value and enter your result below.

```{webr}
#| exercise: prob2-10
# Enter your answer / result here

```

```{webr}
#| exercise: prob2-10
#| solution: true

0
```

```{webr}
#| exercise: prob2-10
#| check: true

exercise_results$prob2_10 <<- "fail"
gradethis::grade_this({
  if (.result == 0) {
    exercise_results$prob2_10 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the standard error and enter your result below.

```{webr}
#| exercise: prob2-11
# Enter your answer / result here

```

```{webr}
#| exercise: prob2-11
#| solution: true

sqrt((4.5^2 / 14) + (6.2^2 / 9))
```

```{webr}
#| exercise: prob2-11
#| check: true

exercise_results$prob2_11 <<- "fail"
gradethis::grade_this({
  if (abs(.result - sqrt((4.5^2/14) + (6.2^2/9))) < 0.01) {
    exercise_results$prob2_11 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the test statistic and enter your result below.

```{webr}
#| exercise: prob2-12
# Enter your answer / result here

```

```{webr}
#| exercise: prob2-12
#| solution: true

(39 - 44) / sqrt((4.5^2 / 14) + (6.2^2 / 9))
```

```{webr}
#| exercise: prob2-12
#| check: true

exercise_results$prob2_12 <<- "fail"
gradethis::grade_this({
  if (abs(abs(.result) - abs((39 - 44) / sqrt((4.5^2/14) + (6.2^2/9)))) < 0.01) {
    exercise_results$prob2_12 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the $p$-value and enter your result below.

```{webr}
#| exercise: prob2-13
# Enter your answer / result here

```

```{webr}
#| exercise: prob2-13
#| solution: true

ts <- (39 - 44) / sqrt((4.5^2 / 14) + (6.2^2 / 9))
pt(-abs(ts), df = 8)
```

```{webr}
#| exercise: prob2-13
#| check: true

exercise_results$prob2_13 <<- "fail"
gradethis::grade_this({
  ts <- (39 - 44) / sqrt((4.5^2/14) + (6.2^2/9))
  if (abs(.result - pt(-abs(ts), df = 8)) < 0.0005) {
    exercise_results$prob2_13 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-caution}
## Problem 2, Part X

What is the result of the test?

```{ojs}
//| echo: false

viewof q15 = Inputs.radio(
  new Map([
    ["Since p ≥ α, we do not have enough evidence to reject the null hypothesis.", 1],
    ["Since p ≥ α, we accept the null hypothesis.", 2],
    ["Since p < α, we reject the null hypothesis and accept the alternative hypothesis.", 3],
    ["Since p < α, we fail to reject the null hypothesis.", 4],
    ["It is impossible to determine.", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q15_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q15_selected_t16ae", JSON.stringify(q15));
  localStorage.setItem("q15_correct_t16ae", "3");
  localStorage.setItem("q15_result_t16ae", q15 === null ? "unattempted" : (q15 == 3 ? "correct" : "incorrect"));
}
ok_response(q15, "3");
```

:::

:::{.callout-caution}
## Problem 2, Part XI

The result of the test means that:

```{ojs}
//| echo: false

viewof q16 = Inputs.radio(
  new Map([
    ["The sample data did not provide significant evidence to suggest that the average age of men on Death Row is younger than the average age of women on Death Row.", 1],
    ["The sample data proved that the average age of men on Death Row is not younger than the average age of women on Death Row.", 2],
    ["The sample data provided significant evidence to suggest that the average age of men on Death Row is younger than the average age of women on Death Row.", 3],
    ["The sample data provided significant evidence to suggest that the sample average age of men on Death Row is younger than the sample average age of women on Death Row.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q16_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q16_selected_t16ae", JSON.stringify(q16));
  localStorage.setItem("q16_correct_t16ae", "3");
  localStorage.setItem("q16_result_t16ae", q16 === null ? "unattempted" : (q16 == 3 ? "correct" : "incorrect"));
}
ok_response(q16, "3");
```

:::


---

## Problem 3

A recent random sample of 30 customers who switched to a car insurance company boasting that new customers save an average of \$534 per year resulted in an average savings of \$525 with a standard deviation of \$40. Find a 90% confidence interval for the true mean savings of customers switching to this company. ([See worked solution](www/SampleInferenceCarInsurance_Spreadsheets.pdf){target="_blank"})

:::{.callout-caution}
## Problem 3, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

viewof q17 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q17_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q17_selected_t16ae", JSON.stringify(q17));
  localStorage.setItem("q17_correct_t16ae", "2");
  localStorage.setItem("q17_result_t16ae", q17 === null ? "unattempted" : (q17 == 2 ? "correct" : "incorrect"));
}
ok_response(q17, "2");
```

:::

:::{.callout-caution}
## Problem 3, Part II

What is the desired level of confidence?

```{ojs}
//| echo: false

viewof q18 = Inputs.radio(
  new Map([
    ["90%", 1],
    ["95%", 2],
    ["98%", 3],
    ["99%", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q18_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q18_selected_t16ae", JSON.stringify(q18));
  localStorage.setItem("q18_correct_t16ae", "1");
  localStorage.setItem("q18_result_t16ae", q18 === null ? "unattempted" : (q18 == 1 ? "correct" : "incorrect"));
}
ok_response(q18, "1");
```

:::

:::{.callout-caution}
## Problem 3, Part III

Is your confidence interval being built to capture a mean ($\mu$), a proportion ($p$), or something else?

```{ojs}
//| echo: false

viewof q19 = Inputs.radio(
  new Map([
    ["A mean or difference in means.", 1],
    ["A proportion or difference in proportions.", 2],
    ["Something else altogether.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q19_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q19_selected_t16ae", JSON.stringify(q19));
  localStorage.setItem("q19_correct_t16ae", "1");
  localStorage.setItem("q19_result_t16ae", q19 === null ? "unattempted" : (q19 == 1 ? "correct" : "incorrect"));
}
ok_response(q19, "1");
```

:::

:::{.callout-caution}
## Problem 3, Part IV

Does the population parameter belong to a single group or is it a comparison of multiple groups?

```{ojs}
//| echo: false

viewof q20 = Inputs.radio(
  new Map([
    ["The population parameter belongs to a single group.", 1],
    ["The population parameter involves a comparison of two groups.", 2],
    ["The population parameter involves a comparison of more than two groups.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q20_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q20_selected_t16ae", JSON.stringify(q20));
  localStorage.setItem("q20_correct_t16ae", "1");
  localStorage.setItem("q20_result_t16ae", q20 === null ? "unattempted" : (q20 == 1 ? "correct" : "incorrect"));
}
ok_response(q20, "1");
```

:::

:::{.callout-caution}
## Problem 3, Part V

Do we know the population standard deviation ($\sigma$) for dollars saved?

```{ojs}
//| echo: false

viewof q21 = Inputs.radio(
  new Map([
    ["Yes, the population standard deviation is known.", 1],
    ["No, the population standard deviation is not known.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q21_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q21_selected_t16ae", JSON.stringify(q21));
  localStorage.setItem("q21_correct_t16ae", "2");
  localStorage.setItem("q21_result_t16ae", q21 === null ? "unattempted" : (q21 == 2 ? "correct" : "incorrect"));
}
ok_response(q21, "2");
```

:::

:::{.callout-caution}
## Problem 3, Part VI

Which standard error formula should be used?

```{ojs}
//| echo: false

viewof q22 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q22_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q22_selected_t16ae", JSON.stringify(q22));
  localStorage.setItem("q22_correct_t16ae", "2");
  localStorage.setItem("q22_result_t16ae", q22 === null ? "unattempted" : (q22 == 2 ? "correct" : "incorrect"));
}
ok_response(q22, "2");
```

:::

:::{.callout-caution}
## Problem 3, Part VII

Which distribution should be used to identify the critical value?

```{ojs}
//| echo: false

viewof q23 = Inputs.radio(
  new Map([
    ["The normal distribution.", 1],
    ["The t-distribution with df = n - 1.", 2],
    ["The t-distribution with df = min(n1-1, n2-1).", 3],
    ["The t-distribution with df = n_diff - 1.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q23_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q23_selected_t16ae", JSON.stringify(q23));
  localStorage.setItem("q23_correct_t16ae", "2");
  localStorage.setItem("q23_result_t16ae", q23 === null ? "unattempted" : (q23 == 2 ? "correct" : "incorrect"));
}
ok_response(q23, "2");
```

:::

Use your spreadsheet to compute the critical value and enter your result below.

```{webr}
#| exercise: prob3-7
# Enter your answer / result here

```

```{webr}
#| exercise: prob3-7
#| solution: true

qt(0.95, df = 29)
```

```{webr}
#| exercise: prob3-7
#| check: true

exercise_results$prob3_7 <<- "fail"
gradethis::grade_this({
  if (abs(abs(.result) - qt(0.95, df = 29)) < 0.005) {
    exercise_results$prob3_7 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the point estimate and enter your result below.

```{webr}
#| exercise: prob3-8
# Enter your answer / result here

```

```{webr}
#| exercise: prob3-8
#| solution: true

525
```

```{webr}
#| exercise: prob3-8
#| check: true

exercise_results$prob3_8 <<- "fail"
gradethis::grade_this({
  if (.result == 525) {
    exercise_results$prob3_8 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the standard error and enter your result below.

```{webr}
#| exercise: prob3-9
# Enter your answer / result here

```

```{webr}
#| exercise: prob3-9
#| solution: true

40 / sqrt(30)
```

```{webr}
#| exercise: prob3-9
#| check: true

exercise_results$prob3_9 <<- "fail"
gradethis::grade_this({
  if (abs(.result - (40 / sqrt(30))) < 0.01) {
    exercise_results$prob3_9 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the lower bound of the confidence interval and enter your result below.

```{webr}
#| exercise: prob3-10
# Enter your answer / result here

```

```{webr}
#| exercise: prob3-10
#| solution: true

525 - qt(0.95, df = 29) * (40 / sqrt(30))
```

```{webr}
#| exercise: prob3-10
#| check: true

exercise_results$prob3_10 <<- "fail"
gradethis::grade_this({
  if (abs(.result - (525 - qt(0.95, df = 29) * 40 / sqrt(30))) < 0.01) {
    exercise_results$prob3_10 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the upper bound of the confidence interval and enter your result below.

```{webr}
#| exercise: prob3-11
# Enter your answer / result here

```

```{webr}
#| exercise: prob3-11
#| solution: true

525 + qt(0.95, df = 29) * (40 / sqrt(30))
```

```{webr}
#| exercise: prob3-11
#| check: true

exercise_results$prob3_11 <<- "fail"
gradethis::grade_this({
  if (abs(.result - (525 + qt(0.95, df = 29) * 40 / sqrt(30))) < 0.01) {
    exercise_results$prob3_11 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-caution}
## Problem 3, Part VIII

The correct interpretation of this confidence interval is:

```{ojs}
//| echo: false

viewof q24 = Inputs.radio(
  new Map([
    ["We are 90% confident that the true average savings customers should expect when switching to this provider is between the lower and upper bounds calculated.", 1],
    ["90% of the savings amounts customers realize when switching to this provider will fall between the lower and upper bounds calculated.", 2],
    ["We are 90% confident that the average savings customers in this sample experienced when switching to this provider fall between the lower and upper bounds calculated.", 3],
    ["There is a 90% chance a customer switching to this provider will realize savings between the lower and upper bounds calculated.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q24_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q24_selected_t16ae", JSON.stringify(q24));
  localStorage.setItem("q24_correct_t16ae", "1");
  localStorage.setItem("q24_result_t16ae", q24 === null ? "unattempted" : (q24 == 1 ? "correct" : "incorrect"));
}
ok_response(q24, "1");
```

:::


---

## Problem 4

In January 2011, The Marist Poll published a report stating that 66% of adults nationally think licensed drivers should be required to retake their road test once they reach 65 years of age. In that same year, 200 random citizens of New Hampshire were asked whether they were in favor of an additional road test for drivers at 65 years of age — 118 responded that they were in favor. Construct a 95% confidence interval for the true proportion of New Hampshirites who are in favor of this proposal. ([See worked solution](www/SampleInferenceMaristPoll_Spreadsheets.pdf){target="_blank"})

:::{.callout-caution}
## Problem 4, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

viewof q25 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q25_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q25_selected_t16ae", JSON.stringify(q25));
  localStorage.setItem("q25_correct_t16ae", "2");
  localStorage.setItem("q25_result_t16ae", q25 === null ? "unattempted" : (q25 == 2 ? "correct" : "incorrect"));
}
ok_response(q25, "2");
```

:::

:::{.callout-caution}
## Problem 4, Part II

What is the desired level of confidence?

```{ojs}
//| echo: false

viewof q26 = Inputs.radio(
  new Map([
    ["90%", 1],
    ["95%", 2],
    ["98%", 3],
    ["99%", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q26_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q26_selected_t16ae", JSON.stringify(q26));
  localStorage.setItem("q26_correct_t16ae", "2");
  localStorage.setItem("q26_result_t16ae", q26 === null ? "unattempted" : (q26 == 2 ? "correct" : "incorrect"));
}
ok_response(q26, "2");
```

:::

:::{.callout-caution}
## Problem 4, Part III

Is your confidence interval being built to capture a mean ($\mu$), a proportion ($p$), or something else?

```{ojs}
//| echo: false

viewof q27 = Inputs.radio(
  new Map([
    ["A mean or difference in means.", 1],
    ["A proportion or difference in proportions.", 2],
    ["Something else altogether.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q27_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q27_selected_t16ae", JSON.stringify(q27));
  localStorage.setItem("q27_correct_t16ae", "2");
  localStorage.setItem("q27_result_t16ae", q27 === null ? "unattempted" : (q27 == 2 ? "correct" : "incorrect"));
}
ok_response(q27, "2");
```

:::

:::{.callout-caution}
## Problem 4, Part IV

Does the population parameter belong to a single group or is it a comparison of multiple groups?

```{ojs}
//| echo: false

viewof q28 = Inputs.radio(
  new Map([
    ["The population parameter belongs to a single group.", 1],
    ["The population parameter involves a comparison of two groups.", 2],
    ["The population parameter involves a comparison of more than two groups.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q28_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q28_selected_t16ae", JSON.stringify(q28));
  localStorage.setItem("q28_correct_t16ae", "1");
  localStorage.setItem("q28_result_t16ae", q28 === null ? "unattempted" : (q28 == 1 ? "correct" : "incorrect"));
}
ok_response(q28, "1");
```

:::

:::{.callout-caution}
## Problem 4, Part V

Which standard error formula should be used?

```{ojs}
//| echo: false

viewof q29 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q29_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q29_selected_t16ae", JSON.stringify(q29));
  localStorage.setItem("q29_correct_t16ae", "4");
  localStorage.setItem("q29_result_t16ae", q29 === null ? "unattempted" : (q29 == 4 ? "correct" : "incorrect"));
}
ok_response(q29, "4");
```

:::

:::{.callout-caution}
## Problem 4, Part VI

Which distribution should be used to identify the critical value?

```{ojs}
//| echo: false

viewof q30 = Inputs.radio(
  new Map([
    ["The normal distribution.", 1],
    ["The t-distribution with df = n - 1.", 2],
    ["The t-distribution with df = min(n1-1, n2-1).", 3],
    ["The t-distribution with df = n_diff - 1.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q30_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q30_selected_t16ae", JSON.stringify(q30));
  localStorage.setItem("q30_correct_t16ae", "1");
  localStorage.setItem("q30_result_t16ae", q30 === null ? "unattempted" : (q30 == 1 ? "correct" : "incorrect"));
}
ok_response(q30, "1");
```

:::

Use your spreadsheet to compute the critical value and enter your result below.

```{webr}
#| exercise: prob4-6
# Enter your answer / result here

```

```{webr}
#| exercise: prob4-6
#| solution: true

qnorm(0.975)
```

```{webr}
#| exercise: prob4-6
#| check: true

exercise_results$prob4_6 <<- "fail"
gradethis::grade_this({
  if (abs(abs(.result) - qnorm(0.975)) < 0.005) {
    exercise_results$prob4_6 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the point estimate and enter your result below.

```{webr}
#| exercise: prob4-7
# Enter your answer / result here

```

```{webr}
#| exercise: prob4-7
#| solution: true

118 / 200
```

```{webr}
#| exercise: prob4-7
#| check: true

exercise_results$prob4_7 <<- "fail"
gradethis::grade_this({
  if (abs(.result - 118/200) < 0.0005) {
    exercise_results$prob4_7 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the standard error and enter your result below.

```{webr}
#| exercise: prob4-8
# Enter your answer / result here

```

```{webr}
#| exercise: prob4-8
#| solution: true

sqrt((118/200) * (1 - 118/200) / 200)
```

```{webr}
#| exercise: prob4-8
#| check: true

exercise_results$prob4_8 <<- "fail"
gradethis::grade_this({
  p_hat <- 118/200
  if (abs(.result - sqrt(p_hat * (1 - p_hat) / 200)) < 0.0005) {
    exercise_results$prob4_8 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the lower bound of the confidence interval and enter your result below.

```{webr}
#| exercise: prob4-9
# Enter your answer / result here

```

```{webr}
#| exercise: prob4-9
#| solution: true

p_hat <- 118/200
p_hat - qnorm(0.975) * sqrt(p_hat * (1 - p_hat) / 200)
```

```{webr}
#| exercise: prob4-9
#| check: true

exercise_results$prob4_9 <<- "fail"
gradethis::grade_this({
  p_hat <- 118/200
  lower <- p_hat - qnorm(0.975) * sqrt(p_hat * (1 - p_hat) / 200)
  if (abs(.result - lower) < 0.0005) {
    exercise_results$prob4_9 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use your spreadsheet to compute the upper bound of the confidence interval and enter your result below.

```{webr}
#| exercise: prob4-10
# Enter your answer / result here

```

```{webr}
#| exercise: prob4-10
#| solution: true

p_hat <- 118/200
p_hat + qnorm(0.975) * sqrt(p_hat * (1 - p_hat) / 200)
```

```{webr}
#| exercise: prob4-10
#| check: true

exercise_results$prob4_10 <<- "fail"
gradethis::grade_this({
  p_hat <- 118/200
  upper <- p_hat + qnorm(0.975) * sqrt(p_hat * (1 - p_hat) / 200)
  if (abs(.result - upper) < 0.0005) {
    exercise_results$prob4_10 <<- "pass"
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-caution}
## Problem 4, Part VII

The correct interpretation of this confidence interval is:

```{ojs}
//| echo: false

viewof q31 = Inputs.radio(
  new Map([
    ["We are 95% confident that the sample proportion of New Hampshirites favoring the policy falls between the lower and upper bounds calculated.", 1],
    ["95% of New Hampshirites favor the policy.", 2],
    ["We are 95% confident that the true proportion of New Hampshirites favoring the policy is between the lower and upper bounds calculated.", 3],
    ["There is a 95% chance the true proportion of New Hampshirites favoring the policy falls between the lower and upper bounds calculated.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q31_selected_t16ae") ?? "null")}
);

{
  localStorage.setItem("q31_selected_t16ae", JSON.stringify(q31));
  localStorage.setItem("q31_correct_t16ae", "3");
  localStorage.setItem("q31_result_t16ae", q31 === null ? "unattempted" : (q31 == 3 ? "correct" : "incorrect"));
}
ok_response(q31, "3");
```

:::


## Submit

If you are part of a course with an instructor who is grading your work on these activities, please copy and submit ***both*** of the hashes below using the method your instructor has requested.

:::{.callout-warning}
## Question Hash

The hash below encodes your responses to the multiple choice questions in this activity.

```{ojs}
//| echo: false

function buildQuestionResults() {
  return {
    notebook: "Topic 16: Inference Practice (Part A, Spreadsheets)",
    type: "questions",
    timestamp: new Date().toISOString(),
    questions: {
      q1_prob1_inference_type: { selected: q1, correct_answer: "1", result: q1 === null ? "unattempted" : (q1 == 1 ? "correct" : "incorrect") },
      q2_prob1_num_groups: { selected: q2, correct_answer: "1", result: q2 === null ? "unattempted" : (q2 == 1 ? "correct" : "incorrect") },
      q3_prob1_pop_sd_known: { selected: q3, correct_answer: "1", result: q3 === null ? "unattempted" : (q3 == 1 ? "correct" : "incorrect") },
      q4_prob1_se_formula: { selected: q4, correct_answer: "1", result: q4 === null ? "unattempted" : (q4 == 1 ? "correct" : "incorrect") },
      q5_prob1_distribution: { selected: q5, correct_answer: "1", result: q5 === null ? "unattempted" : (q5 == 1 ? "correct" : "incorrect") },
      q6_prob2_inference_type: { selected: q6, correct_answer: "3", result: q6 === null ? "unattempted" : (q6 == 3 ? "correct" : "incorrect") },
      q7_prob2_significance_level: { selected: q7, correct_answer: "2", result: q7 === null ? "unattempted" : (q7 == 2 ? "correct" : "incorrect") },
      q8_prob2_parameter_type: { selected: q8, correct_answer: "1", result: q8 === null ? "unattempted" : (q8 == 1 ? "correct" : "incorrect") },
      q9_prob2_num_groups: { selected: q9, correct_answer: "2", result: q9 === null ? "unattempted" : (q9 == 2 ? "correct" : "incorrect") },
      q10_prob2_hypotheses: { selected: q10, correct_answer: "3", result: q10 === null ? "unattempted" : (q10 == 3 ? "correct" : "incorrect") },
      q11_prob2_pop_sd_known: { selected: q11, correct_answer: "2", result: q11 === null ? "unattempted" : (q11 == 2 ? "correct" : "incorrect") },
      q12_prob2_paired: { selected: q12, correct_answer: "1", result: q12 === null ? "unattempted" : (q12 == 1 ? "correct" : "incorrect") },
      q13_prob2_se_formula: { selected: q13, correct_answer: "3", result: q13 === null ? "unattempted" : (q13 == 3 ? "correct" : "incorrect") },
      q14_prob2_distribution: { selected: q14, correct_answer: "3", result: q14 === null ? "unattempted" : (q14 == 3 ? "correct" : "incorrect") },
      q15_prob2_test_result: { selected: q15, correct_answer: "3", result: q15 === null ? "unattempted" : (q15 == 3 ? "correct" : "incorrect") },
      q16_prob2_conclusion: { selected: q16, correct_answer: "3", result: q16 === null ? "unattempted" : (q16 == 3 ? "correct" : "incorrect") },
      q17_prob3_inference_type: { selected: q17, correct_answer: "2", result: q17 === null ? "unattempted" : (q17 == 2 ? "correct" : "incorrect") },
      q18_prob3_confidence_level: { selected: q18, correct_answer: "1", result: q18 === null ? "unattempted" : (q18 == 1 ? "correct" : "incorrect") },
      q19_prob3_parameter_type: { selected: q19, correct_answer: "1", result: q19 === null ? "unattempted" : (q19 == 1 ? "correct" : "incorrect") },
      q20_prob3_num_groups: { selected: q20, correct_answer: "1", result: q20 === null ? "unattempted" : (q20 == 1 ? "correct" : "incorrect") },
      q21_prob3_pop_sd_known: { selected: q21, correct_answer: "2", result: q21 === null ? "unattempted" : (q21 == 2 ? "correct" : "incorrect") },
      q22_prob3_se_formula: { selected: q22, correct_answer: "2", result: q22 === null ? "unattempted" : (q22 == 2 ? "correct" : "incorrect") },
      q23_prob3_distribution: { selected: q23, correct_answer: "2", result: q23 === null ? "unattempted" : (q23 == 2 ? "correct" : "incorrect") },
      q24_prob3_interpretation: { selected: q24, correct_answer: "1", result: q24 === null ? "unattempted" : (q24 == 1 ? "correct" : "incorrect") },
      q25_prob4_inference_type: { selected: q25, correct_answer: "2", result: q25 === null ? "unattempted" : (q25 == 2 ? "correct" : "incorrect") },
      q26_prob4_confidence_level: { selected: q26, correct_answer: "2", result: q26 === null ? "unattempted" : (q26 == 2 ? "correct" : "incorrect") },
      q27_prob4_parameter_type: { selected: q27, correct_answer: "2", result: q27 === null ? "unattempted" : (q27 == 2 ? "correct" : "incorrect") },
      q28_prob4_num_groups: { selected: q28, correct_answer: "1", result: q28 === null ? "unattempted" : (q28 == 1 ? "correct" : "incorrect") },
      q29_prob4_se_formula: { selected: q29, correct_answer: "4", result: q29 === null ? "unattempted" : (q29 == 4 ? "correct" : "incorrect") },
      q30_prob4_distribution: { selected: q30, correct_answer: "1", result: q30 === null ? "unattempted" : (q30 == 1 ? "correct" : "incorrect") },
      q31_prob4_interpretation: { selected: q31, correct_answer: "3", result: q31 === null ? "unattempted" : (q31 == 3 ? "correct" : "incorrect") }
    }
  };
}

function toBase64(str) {
  return btoa(unescape(encodeURIComponent(str)));
}

question_hash = {
  q1; q2; q3; q4; q5; q6; q7; q8; q9; q10; q11; q12; q13; q14; q15; q16;
  q17; q18; q19; q20; q21; q22; q23; q24; q25; q26; q27; q28; q29; q30; q31;
  return toBase64(JSON.stringify(buildQuestionResults()));
}
```

```{ojs}
//| echo: false

html`<div style="font-family: monospace; font-size: 0.85em; background: #f5f5f5; padding: 12px; border-radius: 6px; word-break: break-all; border: 1px solid #ddd; user-select: all; cursor: pointer;" onclick="navigator.clipboard.writeText(this.innerText)">
  ${question_hash}
</div>
<p style="margin-top: 8px; font-size: 0.9em; color: #555;">
  Click the box to copy to clipboard.
</p>`
```

:::

:::{.callout-warning}
## Exercise Hash

Click the button below to generate your exercise submission code. This hash encodes your work on the graded exercises in this activity.

You must have attempted the graded exercises before clicking — clicking generates a snapshot of your current results. If you have completed the activity over multiple sessions, please go back through and enter your answer in each graded exercise cell before generating the hash below, to ensure your most recent results are recorded.

```{webr}
#| exercise: generate-exercise-hash
#| caption: "Generate Exercise Code"

exercise_results$timestamp <- format(Sys.time(), "%Y-%m-%dT%H:%M:%SZ", tz = "UTC")
exercise_results$type <- "exercises"

hash <- base64encode(charToRaw(toJSON(exercise_results, auto_unbox = TRUE)))
cat(hash)
```

:::

## Summary

:::{.callout-tip}
## Main Takeaways

These four problems covered the range of inference tasks you'll encounter most frequently. A few reminders worth carrying forward:

+ **Always identify the task first** — probability, confidence interval, hypothesis test, or sample size. The task determines everything that follows.
+ **The decision tree is your guide** for identifying the correct standard error formula and distribution. The key questions: means or proportions? one group or two? population standard deviation known or unknown?
+ **When $\sigma$ is unknown**, use $s$ in the standard error formula and the $t$-distribution for critical values and $p$-values. When $\sigma$ is known, use the normal distribution.
+ **Interpretation matters.** A confidence interval captures a population parameter with a stated level of confidence. It is not a probability statement about the parameter, and not a statement about individual observations.
:::

:::{.callout-tip}
## Looking Ahead

Part B continues with four more mixed inference problems. The scaffolding remains, but the worked solution links are removed — you're increasingly on your own. Take stock of what is difficult for you and be sure to ask a teacher, mentor, or friend when you need help. Keep the decision tree and general strategy documents close at hand. You'll find them to be helpful.
:::
