---
title: "Topic 16: Inference Practice (Part B)"
format: 
  live-html:
    include-in-header:
      text: |
        <style>
        .btn-exercise-solution { display: none !important; }
        .exercise-editor-body { max-height: 300px; overflow-y: auto; }
        </style>
    webr:
      packages:
        - dplyr
        - ggplot2
        - gradethis
      cell-options:
        persist: true
    live:
      show-solutions: true
---

{{< include ./_extensions/r-wasm/live/_knitr.qmd >}}
{{< include ./_extensions/r-wasm/live/_gradethis.qmd >}}

```{webr}
#| edit: false
#| output: false
#| define:
#|   - ok_response

library(htmltools)
ok_response <- function(response, n) {
  if (is.na(response)) div(HTML("You haven't answered yet."), style = "color: purple")
  else if (response == n) div(HTML("Correct ✓"), style = "color: green")
  else div(HTML("Incorrect ✗"), style = "color: red")
}
```

```{webr}
#| context: setup
#| echo: false
#| output: false

library(dplyr)
library(ggplot2)

theme_set(theme_bw(base_size = 12))
```

```{r setup, include=FALSE}
library(dplyr)
library(ggplot2)

theme_set(theme_bw(base_size = 12))
```

:::{.callout-note}
## About

This is a continuation of the Topic 16 (Part A) activity. This second of three practice notebooks contains four more scaffolded inference problems — Problems 5 through 8. No new content is introduced.
:::

## Inference Practice

This activity continues where Part A left off. For convenience, the reference documents are linked here:

+ [*Standard Error Decision Tree*](https://agmath.github.io/SiteFiles/StdErrorDecisionTree.pdf){target="_blank"}
+ [*General Strategy for Confidence Intervals*](https://agmath.github.io/SiteFiles/CIgeneral.pdf){target="_blank"}
+ [*General Strategy for Hypothesis Tests*](https://agmath.github.io/SiteFiles/HTgeneral.pdf){target="_blank"}

For any problems involving hypothesis tests, assume $\alpha = 0.05$ unless otherwise stated.

:::{.callout-important}
## No Hints

There are no hints included in this activity. Use the reference documents linked above and your previous notebooks to guide you.
:::

---

## Problem 5

A pollster working for Pete Buttigieg is interested in identifying the proportion of registered Democrats planning to vote in the NH primary for whom Mayor Pete is their first or second choice candidate. If the pollster wants to estimate this proportion to within one percentage point at 95% confidence, how many voters should be included in a sample if data from a focus group suggests that the proportion is near 42%?

:::{.callout-caution}
## Problem 5, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

mutable ok_response = (response, n) => { return html`Loading...` };
viewof q1 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q1") ?? "null")}
);

localStorage.setItem("q1", JSON.stringify(q1));
ok_response(q1, "4");
```

:::

:::{.callout-caution}
## Problem 5, Part II

The parameter the pollster is attempting to measure is a:

```{ojs}
//| echo: false

viewof q2 = Inputs.radio(
  new Map([
    ["Population mean.", 1],
    ["Population proportion.", 2],
    ["Sample mean.", 3],
    ["Sample proportion.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q2") ?? "null")}
);

localStorage.setItem("q2", JSON.stringify(q2));
ok_response(q2, "2");
```

:::

Use the code block below to input or compute the desired margin of error.

```{webr}
#| exercise: prob5-3
______
```

```{webr}
#| exercise: prob5-3
#| solution: true

0.01
```

```{webr}
#| exercise: prob5-3
#| check: true

gradethis::grade_this({
  if (.result == 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the critical value ($z_{\alpha/2}$) for a 95% confidence level.

```{webr}
#| exercise: prob5-4
______
```

```{webr}
#| exercise: prob5-4
#| solution: true

qnorm(0.975)
```

```{webr}
#| exercise: prob5-4
#| check: true

gradethis::grade_this({
  if (abs(.result - qnorm(0.975)) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to input the value of $p$ to be used in the sample size calculation.

```{webr}
#| exercise: prob5-5
______
```

```{webr}
#| exercise: prob5-5
#| solution: true

0.42
```

```{webr}
#| exercise: prob5-5
#| check: true

gradethis::grade_this({
  if (.result == 0.42) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the minimum required sample size.

```{webr}
#| exercise: prob5-6
______
```

```{webr}
#| exercise: prob5-6
#| solution: true

ceiling((qnorm(0.975) / 0.01)^2 * 0.42 * (1 - 0.42))
```

```{webr}
#| exercise: prob5-6
#| check: true

gradethis::grade_this({
  required <- ceiling((qnorm(0.975) / 0.01)^2 * 0.42 * (1 - 0.42))
  if (abs(.result - required) < 10 && .result == ceiling(.result)) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-tip}
## Reminder: Always Round Up

Did you remember to round your answer up to the nearest whole number? Rounding down results in either a loss of confidence or an increased margin of error, violating the stated requirements for the estimate.
:::

---

## Problem 6

A Gallup poll surveyed 1,000 randomly sampled Greeks in 2011 and found that 25% of them said they would rate their lives poorly enough to be considered suffering. Construct a 95% confidence interval for the true proportion of Greeks who are suffering.

:::{.callout-caution}
## Problem 6, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

viewof q3 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q3") ?? "null")}
);

localStorage.setItem("q3", JSON.stringify(q3));
ok_response(q3, "2");
```

:::

:::{.callout-caution}
## Problem 6, Part II

What is the desired level of confidence?

```{ojs}
//| echo: false

viewof q4 = Inputs.radio(
  new Map([
    ["90%", 1],
    ["95%", 2],
    ["98%", 3],
    ["99%", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q4") ?? "null")}
);

localStorage.setItem("q4", JSON.stringify(q4));
ok_response(q4, "2");
```

:::

:::{.callout-caution}
## Problem 6, Part III

Is your confidence interval being built to capture a mean ($\mu$), a proportion ($p$), or something else?

```{ojs}
//| echo: false

viewof q5 = Inputs.radio(
  new Map([
    ["A mean or difference in means.", 1],
    ["A proportion or difference in proportions.", 2],
    ["Something else altogether.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q5") ?? "null")}
);

localStorage.setItem("q5", JSON.stringify(q5));
ok_response(q5, "2");
```

:::

:::{.callout-caution}
## Problem 6, Part IV

Does the population parameter belong to a single group or is it a comparison of multiple groups?

```{ojs}
//| echo: false

viewof q6 = Inputs.radio(
  new Map([
    ["The population parameter belongs to a single group.", 1],
    ["The population parameter involves a comparison of two groups.", 2],
    ["The population parameter involves a comparison of more than two groups.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q6") ?? "null")}
);

localStorage.setItem("q6", JSON.stringify(q6));
ok_response(q6, "1");
```

:::

:::{.callout-caution}
## Problem 6, Part V

Which standard error formula should be used?

```{ojs}
//| echo: false

viewof q7 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q7") ?? "null")}
);

localStorage.setItem("q7", JSON.stringify(q7));
ok_response(q7, "4");
```

:::

:::{.callout-caution}
## Problem 6, Part VI

Which distribution should be used to identify the critical value?

```{ojs}
//| echo: false

viewof q8 = Inputs.radio(
  new Map([
    ["The normal distribution.", 1],
    ["The t-distribution with df = n - 1.", 2],
    ["The t-distribution with df = min(n1-1, n2-1).", 3],
    ["The t-distribution with df = n_diff - 1.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q8") ?? "null")}
);

localStorage.setItem("q8", JSON.stringify(q8));
ok_response(q8, "1");
```

:::

Use the code block below to compute the critical value.

```{webr}
#| exercise: prob6-6
______
```

```{webr}
#| exercise: prob6-6
#| solution: true

qnorm(0.975)
```

```{webr}
#| exercise: prob6-6
#| check: true

gradethis::grade_this({
  if (abs(abs(.result) - qnorm(0.975)) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the point estimate.

```{webr}
#| exercise: prob6-7
______
```

```{webr}
#| exercise: prob6-7
#| solution: true

0.25
```

```{webr}
#| exercise: prob6-7
#| check: true

gradethis::grade_this({
  if (.result == 0.25) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the standard error.

```{webr}
#| exercise: prob6-8
______
```

```{webr}
#| exercise: prob6-8
#| solution: true

sqrt(0.25 * (1 - 0.25) / 1000)
```

```{webr}
#| exercise: prob6-8
#| check: true

gradethis::grade_this({
  if (abs(.result - sqrt(0.25 * 0.75 / 1000)) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the lower bound of the confidence interval.

```{webr}
#| exercise: prob6-9
______
```

```{webr}
#| exercise: prob6-9
#| solution: true

0.25 - qnorm(0.975) * sqrt(0.25 * 0.75 / 1000)
```

```{webr}
#| exercise: prob6-9
#| check: true

gradethis::grade_this({
  if (abs(.result - (0.25 - qnorm(0.975) * sqrt(0.25 * 0.75 / 1000))) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the upper bound of the confidence interval.

```{webr}
#| exercise: prob6-10
______
```

```{webr}
#| exercise: prob6-10
#| solution: true

0.25 + qnorm(0.975) * sqrt(0.25 * 0.75 / 1000)
```

```{webr}
#| exercise: prob6-10
#| check: true

gradethis::grade_this({
  if (abs(.result - (0.25 + qnorm(0.975) * sqrt(0.25 * 0.75 / 1000))) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-caution}
## Problem 6, Part VII

The correct interpretation of this confidence interval is:

```{ojs}
//| echo: false

viewof q9 = Inputs.radio(
  new Map([
    ["There is a 95% chance the true proportion of Greeks who would describe their quality of life as poor enough to be considered 'suffering' falls between the lower and upper bounds calculated.", 1],
    ["95% of Greeks would describe their quality of life as poor enough to be considered 'suffering'.", 2],
    ["We are 95% confident that the sample proportion of Greeks who would describe their quality of life as poor enough to be considered 'suffering' falls between the lower and upper bounds calculated.", 3],
    ["We are 95% confident that the true proportion of Greeks who would describe their quality of life as poor enough to be considered 'suffering' is between the lower and upper bounds calculated.", 4]
    
  ]),
  {value: JSON.parse(localStorage.getItem("q9") ?? "null")}
);

localStorage.setItem("q9", JSON.stringify(q9));
ok_response(q9, "4");
```

:::

---

## Problem 7

A "social experiment" conducted by a TV program questioned what people do when they see a very angry father upset with his child about their poor athletic performance. On two different occasions at the same restaurant, a father and child were depicted — in one scenario the child was a son and in the other a daughter. The table below shows how many restaurant diners were present under each scenario and whether or not they intervened.

| | Son | Daughter | Total |
|:---|:---:|:---:|:---:|
| **Intervened** | 15 | 25 | 40 |
| **No Action** | 25 | 20 | 45 |
| **Total** | 40 | 45 | 85 |

Conduct a hypothesis test to determine whether there is evidence to suggest that the proportion of individuals willing to intervene is dependent on the gender of the child.

:::{.callout-caution}
## Problem 7, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

viewof q10 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q10") ?? "null")}
);

localStorage.setItem("q10", JSON.stringify(q10));
ok_response(q10, "3");
```

:::

:::{.callout-caution}
## Problem 7, Part II

What is the level of significance associated with this test?

```{ojs}
//| echo: false

viewof q11 = Inputs.radio(
  new Map([
    ["α = 0.01", 1],
    ["α = 0.05", 2],
    ["α = 0.10", 3],
    ["The p-value.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q11") ?? "null")}
);

localStorage.setItem("q11", JSON.stringify(q11));
ok_response(q11, "2");
```

:::

:::{.callout-caution}
## Problem 7, Part III

Does this hypothesis test involve testing a statement about a mean ($\mu$), a proportion ($p$), or something else?

```{ojs}
//| echo: false

viewof q12 = Inputs.radio(
  new Map([
    ["One or more means.", 1],
    ["One or more proportions.", 2],
    ["Something else altogether.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q12") ?? "null")}
);

localStorage.setItem("q12", JSON.stringify(q12));
ok_response(q12, "2");
```

:::

:::{.callout-caution}
## Problem 7, Part IV

How many groups are being compared in this test?

```{ojs}
//| echo: false

viewof q13 = Inputs.radio(
  new Map([
    ["The test involves only a single group.", 1],
    ["The test compares two groups.", 2],
    ["The test compares more than two groups.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q13") ?? "null")}
);

localStorage.setItem("q13", JSON.stringify(q13));
ok_response(q13, "2");
```

:::

:::{.callout-caution}
## Problem 7, Part V

Which of the following are the hypotheses associated with this test?

```{ojs}
//| echo: false

viewof q14 = Inputs.radio(
  new Map([
    ["H₀: p_son − p_daughter = 0; Hₐ: p_son − p_daughter ≠ 0", 1],
    ["H₀: p_son − p_daughter = 0; Hₐ: p_son − p_daughter > 0", 2],
    ["H₀: p_son − p_daughter = 0; Hₐ: p_son − p_daughter < 0", 3],
    ["H₀: p_son ≠ p_daughter; Hₐ: p_son = p_daughter", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q14") ?? "null")}
);

localStorage.setItem("q14", JSON.stringify(q14));
ok_response(q14, "1");
```

:::

:::{.callout-caution}
## Problem 7, Part VI

Which standard error formula should be used?

```{ojs}
//| echo: false

viewof q15 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q15") ?? "null")}
);

localStorage.setItem("q15", JSON.stringify(q15));
ok_response(q15, "5");
```

:::

:::{.callout-caution}
## Problem 7, Part VII

Which distribution does the test statistic follow?

```{ojs}
//| echo: false

viewof q16 = Inputs.radio(
  new Map([
    ["The normal distribution.", 1],
    ["The t-distribution with df = n - 1.", 2],
    ["The t-distribution with df = min(n1-1, n2-1).", 3],
    ["The t-distribution with df = n_diff - 1.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q16") ?? "null")}
);

localStorage.setItem("q16", JSON.stringify(q16));
ok_response(q16, "1");
```

:::

Use the code block below to compute the point estimate.

```{webr}
#| exercise: prob7-7
______
```

```{webr}
#| exercise: prob7-7
#| solution: true

(15/40) - (25/45)
```

```{webr}
#| exercise: prob7-7
#| check: true

gradethis::grade_this({
  if (abs(abs(.result) - abs((15/40) - (25/45))) < 0.005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the null value.

```{webr}
#| exercise: prob7-8
______
```

```{webr}
#| exercise: prob7-8
#| solution: true

0
```

```{webr}
#| exercise: prob7-8
#| check: true

gradethis::grade_this({
  if (.result == 0) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the standard error.

```{webr}
#| exercise: prob7-9
______
```

```{webr}
#| exercise: prob7-9
#| solution: true

sqrt(((15/40) * (1 - 15/40) / 40) + ((25/45) * (1 - 25/45) / 45))
```

```{webr}
#| exercise: prob7-9
#| check: true

gradethis::grade_this({
  se <- sqrt(((15/40) * (1 - 15/40) / 40) + ((25/45) * (1 - 25/45) / 45))
  if (abs(.result - se) < 0.0005) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the test statistic.

```{webr}
#| exercise: prob7-10
______
```

```{webr}
#| exercise: prob7-10
#| solution: true

se <- sqrt(((15/40) * (1 - 15/40) / 40) + ((25/45) * (1 - 25/45) / 45))
((15/40) - (25/45)) / se
```

```{webr}
#| exercise: prob7-10
#| check: true

gradethis::grade_this({
  se <- sqrt(((15/40) * (1 - 15/40) / 40) + ((25/45) * (1 - 25/45) / 45))
  ts <- ((15/40) - (25/45)) / se
  if (abs(abs(.result) - abs(ts)) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the $p$-value.

```{webr}
#| exercise: prob7-11
______
```

```{webr}
#| exercise: prob7-11
#| solution: true

se <- sqrt(((15/40) * (1 - 15/40) / 40) + ((25/45) * (1 - 25/45) / 45))
ts <- ((15/40) - (25/45)) / se
2 * (1 - pnorm(abs(ts)))
```

```{webr}
#| exercise: prob7-11
#| check: true

gradethis::grade_this({
  se <- sqrt(((15/40) * (1 - 15/40) / 40) + ((25/45) * (1 - 25/45) / 45))
  ts <- ((15/40) - (25/45)) / se
  pval <- 2 * (1 - pnorm(abs(ts)))
  if (abs(.result - pval) < 0.001) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-caution}
## Problem 7, Part VIII

What is the result of the test?

```{ojs}
//| echo: false

viewof q17 = Inputs.radio(
  new Map([
    ["Since p < α, we reject the null hypothesis and accept the alternative hypothesis.", 1],
    ["Since p < α, we fail to reject the null hypothesis.", 2],
    ["Since p ≥ α, we do not have enough evidence to reject the null hypothesis.", 3],
    ["Since p ≥ α, we accept the null hypothesis.", 4],
    ["It is impossible to determine.", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q17") ?? "null")}
);

localStorage.setItem("q17", JSON.stringify(q17));
ok_response(q17, "3");
```

:::

:::{.callout-caution}
## Problem 7, Part IX

The result of the test means that:

```{ojs}
//| echo: false

viewof q18 = Inputs.radio(
  new Map([
    ["The sample data proved that there is no difference in the proportion of diners willing to intervene based on the gender of the child.", 1],
    ["The sample data did not provide significant evidence to suggest that there is a difference in the proportion of diners willing to intervene based on the gender of the child.", 2],
    ["The sample data provided significant evidence to suggest a difference in the proportion of diners willing to intervene based on the gender of the child.", 3],
    ["The sample data provided significant evidence to suggest a difference in the sample proportion of diners willing to intervene based on the gender of the child.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q18") ?? "null")}
);

localStorage.setItem("q18", JSON.stringify(q18));
ok_response(q18, "2");
```

:::

---

## Problem 8

A University has a policy that its sports teams require no more than 15 hours per week of time commitment from players during the off-season. A random survey of 12 players from the University's Quidditch team during the Fall Semester (Quidditch is a Spring sport) resulted in a mean of 15.25 hours per week of Quidditch-related activity, with a standard deviation of 1 hour. Is there evidence to suggest that the University Quidditch coach is not following University policy? Use $\alpha = 0.10$.

:::{.callout-caution}
## Problem 8, Part I

To answer the question as asked, we should:

```{ojs}
//| echo: false

viewof q19 = Inputs.radio(
  new Map([
    ["Compute a probability.", 1],
    ["Construct a confidence interval.", 2],
    ["Conduct a hypothesis test.", 3],
    ["Find a required sample size.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q19") ?? "null")}
);

localStorage.setItem("q19", JSON.stringify(q19));
ok_response(q19, "3");
```

:::

:::{.callout-caution}
## Problem 8, Part II

What is the level of significance associated with this test?

```{ojs}
//| echo: false

viewof q20 = Inputs.radio(
  new Map([
    ["α = 0.01", 1],
    ["α = 0.05", 2],
    ["α = 0.10", 3],
    ["The p-value.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q20") ?? "null")}
);

localStorage.setItem("q20", JSON.stringify(q20));
ok_response(q20, "3");
```

:::

:::{.callout-caution}
## Problem 8, Part III

Does this hypothesis test involve testing a statement about a mean ($\mu$), a proportion ($p$), or something else?

```{ojs}
//| echo: false

viewof q21 = Inputs.radio(
  new Map([
    ["One or more means.", 1],
    ["One or more proportions.", 2],
    ["Something else altogether.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q21") ?? "null")}
);

localStorage.setItem("q21", JSON.stringify(q21));
ok_response(q21, "1");
```

:::

:::{.callout-caution}
## Problem 8, Part IV

How many groups are being compared in this test?

```{ojs}
//| echo: false

viewof q22 = Inputs.radio(
  new Map([
    ["The test involves only a single group.", 1],
    ["The test compares two groups.", 2],
    ["The test compares more than two groups.", 3]
  ]),
  {value: JSON.parse(localStorage.getItem("q22") ?? "null")}
);

localStorage.setItem("q22", JSON.stringify(q22));
ok_response(q22, "1");
```

:::

:::{.callout-caution}
## Problem 8, Part V

Which of the following are the hypotheses associated with this test?

```{ojs}
//| echo: false

viewof q23 = Inputs.radio(
  new Map([
    ["H₀: μ = 15; Hₐ: μ ≠ 15", 1],
    ["H₀: μ = 15; Hₐ: μ > 15", 2],
    ["H₀: μ = 15; Hₐ: μ < 15", 3],
    ["H₀: μ = 15.25; Hₐ: μ < 15.25", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q23") ?? "null")}
);

localStorage.setItem("q23", JSON.stringify(q23));
ok_response(q23, "2");
```

:::

:::{.callout-caution}
## Problem 8, Part VI

Do we know the population standard deviation ($\sigma$) for hours spent on Quidditch?

```{ojs}
//| echo: false

viewof q24 = Inputs.radio(
  new Map([
    ["Yes, the population standard deviation is known.", 1],
    ["No, the population standard deviation is not known.", 2]
  ]),
  {value: JSON.parse(localStorage.getItem("q24") ?? "null")}
);

localStorage.setItem("q24", JSON.stringify(q24));
ok_response(q24, "2");
```

:::

:::{.callout-caution}
## Problem 8, Part VII

Which standard error formula should be used?

```{ojs}
//| echo: false

viewof q25 = Inputs.radio(
  new Map([
    ["SE = σ/sqrt(n)", 1],
    ["SE = s/sqrt(n)", 2],
    ["SE = sqrt(s1²/n1 + s2²/n2)", 3],
    ["SE = sqrt(p(1-p)/n)", 4],
    ["SE = sqrt(p1(1-p1)/n1 + p2(1-p2)/n2)", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q25") ?? "null")}
);

localStorage.setItem("q25", JSON.stringify(q25));
ok_response(q25, "2");
```

:::

:::{.callout-caution}
## Problem 8, Part VIII

Which distribution does the test statistic follow?

```{ojs}
//| echo: false

viewof q26 = Inputs.radio(
  new Map([
    ["The normal distribution.", 1],
    ["The t-distribution with df = n - 1.", 2],
    ["The t-distribution with df = min(n1-1, n2-1).", 3],
    ["The t-distribution with df = n_diff - 1.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q26") ?? "null")}
);

localStorage.setItem("q26", JSON.stringify(q26));
ok_response(q26, "2");
```

:::

Use the code block below to compute the point estimate.

```{webr}
#| exercise: prob8-8
______
```

```{webr}
#| exercise: prob8-8
#| solution: true

15.25
```

```{webr}
#| exercise: prob8-8
#| check: true

gradethis::grade_this({
  if (abs(.result) == 15.25) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the null value.

```{webr}
#| exercise: prob8-9
______
```

```{webr}
#| exercise: prob8-9
#| solution: true

15
```

```{webr}
#| exercise: prob8-9
#| check: true

gradethis::grade_this({
  if (.result == 15) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the standard error.

```{webr}
#| exercise: prob8-10
______
```

```{webr}
#| exercise: prob8-10
#| solution: true

1 / sqrt(12)
```

```{webr}
#| exercise: prob8-10
#| check: true

gradethis::grade_this({
  if (abs(.result - 1/sqrt(12)) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the test statistic.

```{webr}
#| exercise: prob8-11
______
```

```{webr}
#| exercise: prob8-11
#| solution: true

(15.25 - 15) / (1 / sqrt(12))
```

```{webr}
#| exercise: prob8-11
#| check: true

gradethis::grade_this({
  ts <- (15.25 - 15) / (1 / sqrt(12))
  if (abs(.result - ts) < 0.01) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

Use the code block below to compute the $p$-value.

```{webr}
#| exercise: prob8-12
______
```

```{webr}
#| exercise: prob8-12
#| solution: true

ts <- (15.25 - 15) / (1 / sqrt(12))
1 - pt(ts, df = 11)
```

```{webr}
#| exercise: prob8-12
#| check: true

gradethis::grade_this({
  ts <- (15.25 - 15) / (1 / sqrt(12))
  pval <- 1 - pt(ts, df = 11)
  if (abs(.result - pval) < 0.0015) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

:::{.callout-caution}
## Problem 8, Part IX

What is the result of the test?

```{ojs}
//| echo: false

viewof q27 = Inputs.radio(
  new Map([
    ["Since p ≥ α, we accept the null hypothesis.", 1],
    ["Since p ≥ α, we do not have enough evidence to reject the null hypothesis.", 2],
    ["Since p < α, we reject the null hypothesis and accept the alternative hypothesis.", 3],
    ["Since p < α, we fail to reject the null hypothesis.", 4],
    ["It is impossible to determine.", 5]
  ]),
  {value: JSON.parse(localStorage.getItem("q27") ?? "null")}
);

localStorage.setItem("q27", JSON.stringify(q27));
ok_response(q27, "2");
```

:::

:::{.callout-caution}
## Problem 8, Part X

The result of the test means that:

```{ojs}
//| echo: false

viewof q28 = Inputs.radio(
  new Map([
    ["The sample data proved that the University Quidditch coach is violating the off-season time commitment policy.", 1],
    ["The sample data did not provide significant evidence to suggest that the University Quidditch coach is violating the off-season time commitment policy.", 2],
    ["The sample data proved that the University Quidditch coach is not violating the off-season time commitment policy.", 3],
    ["The sample data provided significant evidence to suggest that the University Quidditch coach is violating the off-season time commitment policy.", 4]
  ]),
  {value: JSON.parse(localStorage.getItem("q28") ?? "null")}
);

localStorage.setItem("q28", JSON.stringify(q28));
ok_response(q28, "2");
```

:::

## Submit

:::{.callout-warning}

Grading function is not currently available for these web-based notebooks. If you would like to utilize grading functionality, please use either the [Posit Cloud pre-built instance](https://posit.cloud/content/6328402){target="_blank"} or [install the `{learnr}` notebooks locally](https://agmath.github.io/IntroductoryStatistics/AccessingInteractiveNotes.html){target="_blank"}.

:::

## Summary

:::{.callout-tip}
## Main Takeaways

Problems 5–8 introduced a few additional scenarios worth remembering:

+ **Sample size planning** requires rounding up — always. Rounding down means either the confidence level or the margin of error requirement is violated.
+ **When no prior estimate of $p$ is available**, use $p = 0.5$ as a conservative worst-case choice. When a prior estimate exists, use it — it gives a more efficient (smaller) required sample size.
+ **Two-proportion tests use the normal distribution** for the test statistic since we are working with sample proportions, not means.
+ **Single-group tests for a mean** use the $t$-distribution when $\sigma$ is unknown — which is almost always.
:::

:::{.callout-tip}
## Looking Ahead

Part C contains the final four problems — Problems 9 through 12. The scaffolding remains but there are no safety nets. You're nearly there!
:::
